<?xml version="1.0" encoding="UTF-8"?>
<!-- Gegenereerd door tools/convert.py uit archief/ ; niet met de hand bewerken. -->
<chapter xml:id="les-18">
  <title>18: Efficiëntie</title>
  <section xml:id="les-18-s0">
    <title>Inleiding</title>
    <p>Veel programmeertaken kunnen op meer dan één manier worden uitgevoerd, maar de ene manier kan veel sneller zijn dan een andere. Snelle programma's ontwerpen is een deel van de kunst en wetenschap van het programmeren van een computer. In deze oefening kijken we naar een aantal voorbeelden.</p>
    <p><alert>Deel 1: Bereken niet twee keer hetzelfde</alert></p>
    <p>De <url href="http://en.wikipedia.org/wiki/Fibonacci_number" visual="en.wikipedia.org/wiki/Fibonacci_number">Rij van Fibonacci</url> is een fascinerende maar ook eenvoudige rij getallen. Begin met twee getallen, 1 en 1. Dan is de regel:<em> om het volgende getal te krijgen tel je de vorige twee op</em>. Het volgende getal is dus 1+1=2. Dan hebben we de eerste drie elementen,</p>
    <p><c>1, 1, 2</c></p>
    <p>en het vierde element is 1+2=3, vervolgens krijgen we 2+3=5, enzovoort:</p>
    <p><c>1, 1, 2, 3, 5, 8, 13, ...</c></p>
    <p>De Rij van Fibonacci werd aanvankelijk gebruikt om de populatie van konijnen te beschrijven, en het heeft ook fantastische relaties met de architectuur van planten. Hier is een deel van een videoserie over de Rij van Fibonacci:</p>
    <video youtube="ahXIMUkSXX0" width="80%" label="les-18-video-1"/>
    <p>De definitie van de de Rij van Fibonacci leent zich natuurlijk voor een recursieve functie. In de volgende oefening wordt een functie <c>Fibonacci(n)</c> gedefinieerd die het <c>n</c>de element levert uit de rij hierboven (te beginnen met <c>n=1</c>).</p>
    <exercise label="oef-18-nibofacci" language="python">
      <title>Nibofacci</title>
      <statement>
        <p>Sleep in dit programma de regels in de juiste volgorde zodat er een recursieve functie ontstaat dat de Rij van Fibonacci oplevert. De grader zal de eerste tien ervan afdrukken.</p>
      </statement>
      <blocks>
        <block>
          <cline>def Fibonacci(n):</cline>
        </block>
        <block>
          <cline>    if (n==1 or n==2):</cline>
        </block>
        <block>
          <cline>        return 1</cline>
        </block>
        <block>
          <cline>    else:</cline>
        </block>
        <block>
          <cline>        return Fibonacci(n-1) + Fibonacci(n-2)</cline>
        </block>
      </blocks>
    </exercise>
    <p>Deze recursieve functie wordt echter te langzaam om elementen te berekenen die een eind verderop in de rij staan. Druk op <alert>Test statements invoeren</alert> en type <c>print(Fibonacci(80))</c>. Wanneer je dit test krijg je "Time Limit Exceeded."</p>
    <p>Waarom gaat dat zo langzaam? De functie kent geen complexe instructies of loops, alleen maar optellen. Dus moet de traagheid wel gerelateerd zijn aan het aantal keren dat de functie wordt aangeroepen. Wanneer we <c>Fibonacci(3)</c> aanroepen, wordt de functie in totaal drie keer aangeroepen: de begin-aanroep, en dan twee recursieve aanroepen. Wanneer we <c>Fibonacci(4) </c>aanroepen, wordt de recursive functie vijf keer aangeroepen: de begin-aanroep de drie keer die we juist noemden bij <c>n=3</c>, en nog een recursieve aanroep met <c>n=2</c>. Bij het berekenen van <c>Fibonacci(5)</c> gebruiken we in totaal negen aanroepen, en bij <c>Fibonacci(6)</c> zijn er 9+5+1=15 aanroepen. Het aantal aanroepen wordt snel erg groot bij toenemende <c>n</c>!</p>
    <p>Als ruwe schatting heeft <c>Fibonacci(n+2)</c> twee maal zoveel aanroepen nodig als <c>Fibonacci(n)</c>, want <c>Fibonacci(n+2) </c>roept <c>Fibonacci(n)</c> eenmaal direct aan, en een ander keer indirect via de recursieve aanroep van <c>Fibonacci(n+1)</c>. Dus is de rekentijd evenredig met met een <em>exponentiële functie</em> tenminste zo groot als (√2)<m>^{\text{n}}</m>. Dit is te langzaam! De aanroep <c>Fibonacci(80)</c> vergt bijvoorbeeld al meer dan 2<m>^{\text{40}}</m> = 1099511627776 recursieve aanroepen.</p>
    <p>Bovenstaand argument schetst het probleem: <c>Fibonacci(n)</c> twee keer aanroepen en het antwoord erop helemaal opnieuw te berekenen is tijdverspilling. We moeten een nieuwe aanpak verzinnen waarbij we geen tijd verspillen met steeds weer hetzelfde keer op keer te berekenen.</p>
  </section>
  <section xml:id="les-18-s1">
    <title>De oplossing</title>
    <p>Laten we proberen iets in Python te schrijven dat meer lijkt op wat we in de inleiding hebben gedaan. We begonnen door 1, 1 op te schrijven. Dan gingen we verder met het uitbreiden van de rij door de laatste twee elementen op te tellen. Hier volgt de code; het is jouw taak om de regels in de juiste volgorde te zetten.</p>
    <exercise label="oef-18-fast-fibonacci" language="python">
      <title>Fast Fibonacci</title>
      <statement>
        <p>Zet de regels, van het programma waarmee snel grote waarden van de Rij van Fibonacci berekend kunnen worden, in de juiste volgorde.</p>
      </statement>
      <blocks>
        <block>
          <cline>def Fibonacci(n):</cline>
        </block>
        <block>
          <cline>    sequence = [0, 1, 1]  # Fibonacci(0) is 0, Fibonacci(1) and Fibonacci(2) are 1</cline>
        </block>
        <block>
          <cline>    for i in range(3, n+1):</cline>
        </block>
        <block>
          <cline>        sequence.append(sequence[i-1] + sequence[i-2])</cline>
        </block>
        <block>
          <cline>    return sequence[n]</cline>
        </block>
      </blocks>
    </exercise>
    <p>Er is steeds nog wat ruimte voor verbeteringen, omdat bij een nieuwe aanroep het niet nodig is de hele array opnieuw te berekenen, maar het is goed genoeg omdat het snel werkt zelfs bij grote waarden van n!</p>
    <p><alert>Deel 2: Bereken geen onnodige dingen</alert></p>
    <p>Ons tweede voorbeeld gaat over de test of een getal <alert>priem </alert>is. Dit is belangrijk voor cryptografie en computerbeveiliging. Een getal is priem wanneer het precies twee delers heeft: 1 en zichzelf. De eerste priemgetallen zijn 2, 3, 5, 7, 11, 13, 17, 19, 23. (21 bijvoorbeeld is niet priem, omdat het, naast 1 en 21, delers 3 en 7 heeft.)</p>
    <p>Hoe kunnen we in Python testen of een getal priem is? We zagen <xref ref="les-07b" text="custom">eerder</xref> hoe we deelbaarheid kunnen testen:</p>
    <p><c>N % D == 0  # will be True if D is a divisor of N, False otherwise</c></p>
    <p>Door te testen op alle mogelijke delers, komen we tot het volgende programma.</p>
    <example>
      <p>Test om na te gaan of een paar getallen priem zijn</p>
      <program interactive="activecode" language="python" label="les-18-vb-2">
        <code>
def isItPrime(N):
  for D in range(2, N):                        # test D from 2 to N-1
    if N % D == 0:                             # is D a divisor of N?
      print(N, "is not prime; divisible by", D)
      return
  print(N, "is prime")                         # there were no divisors
</code>
      </program>
    </example>
    <p>Het werkt! Maar het is te langzaam bij grote getallen. Ga naar <alert>Enter input</alert> en vul in <c>isItPrime(324635459)</c>. Het programma doet er veel te lang over. Wanneer je nog een paar waarden probeert die groter dan 10000000 zijn om na te gaan of die priem zijn, dan zie je dat het programma altijd te weinig tijd heeft, omdat de grader de deelbaarheids-test loop maar 10 miljoen keer per seconde kan uitvoeren. Wanneer we grotere getallen willen testen hebben we een efficiëntere methode nodig. Laten we er voor zorgen dat het programma werkt voor getallen ter grootte van een triljoen (1000000000000)!</p>
    <p>Moeten we echt <alert>alle</alert> getallen testen tussen <c>2</c> en <c>N-1</c>, om na te gaan of <c>N</c> priem is? Hint</p>
    <remark>
      <title>Hint</title>
      <p>Nee. Maar waarom niet? Neem een paar minuten om te kijken of je zelf met een idee kunt komen. Open een kopie van het programma in de console en kijk hoe groot de deler is die het programma doorgaans vindt wanneer <c>N</c> niet priem is.</p>
    </remark>
  </section>
  <section xml:id="les-18-s2">
    <title>Het idee, en een redenering</title>
    <p>Wanneer je de hint hebt gelezen en hebt geëxperimenteerd, dan zou je hebben kunnen opmerken dat wanneer <c>N</c> niet priem is, het programma meestal een betrekkelijk kleine deler vond in vergelijking tot <c>N</c>. Bijvoorbeeld <c>isItPrime(34827948723948723984729834)</c> wordt tamelijk snel uitgevoerd ondanks dat de input enorm groot is, omdat meteen de deler <c>D=2</c> gevonden wordt.</p>
    <p>Misschien is het helemaal niet nodig om alle mogelijke delers te vinden. Is er een grens aan het aantal delers dat we moeten testen, voordat we er zeker van kunnen zeggen dat <c>N</c> priem is? Gelukkig wel! We kunnen beargumenteren dat <em>als </em><c>N</c><em> niet priem is, een van zijn delers ten hoogste </em><c>sqrt(N) </c><em>is.</em> Waarom? Als <c>N</c> niet priem is, dan is er een deler <c>A</c>. Omdat <c>A</c> een deler is, betekent dit dat er een ander getal <c>B</c> is zodat</p>
    <p><c>A*B == N</c></p>
    <p>Laten we de argumentatie voortzetten. Als <c>A &lt;= sqrt(N)</c> of <c>B &lt;= sqrt(N)</c>, zijn we gelukkig: we hebben dan een deler van <c>N</c> die klein genoeg is zoals we wilden. Maar ze vormen ook de enige mogelijkheid: anders krijgen we een tegenspraak</p>
    <p><c>N = A*B &gt; sqrt(N)*sqrt(N) &gt; N</c></p>
    <p>Mooi! Laten we dan nu dit nieuwe idee in Python implementeren. De gemakkelijkste weg is om de oude benadering aan te passen is om een test toe te voegen in de <c>for</c> loop: als <c>D &gt; sqrt(N)</c> (of gelijk daaraan, <c>D*D &gt; N</c>), kunnen we uit de loop stappen en stoppen met testen.</p>
    <example>
      <p>Snellere priem test</p>
      <program interactive="activecode" language="python" label="les-18-vb-3">
        <code>
def isItPrime(N): # same as before
  for D in range(2, N):
    if (D * D &gt; N):          # first added line
      break                  # second added line
    if N % D == 0:
      print(N, "is not prime; divisible by", D)
      return
  print(N, "is prime")

isItPrime(1000006000009)
isItPrime(1666666009999)
</code>
      </program>
    </example>
    <p>Het programma werkt voor op gigantische priemgetallen!</p>
  </section>
  <section xml:id="les-18-s3">
    <title>Slotoefening</title>
    <p>In deze oefening combineren we de priemgetallen uit de tweede helft met de lijst-gebaseerde aanpak uit de eerste helft. Je programma moet een tabel vullen van lengte 1000001 zodanig dat <c>isPrime[N]</c> gelijk wordt aan <c>True </c>als <c>N</c> priem is, en <c>False</c> als <c>N</c> dat niet is, voor alle <c>N</c> tot een miljoen. (<c>isPrime[0]</c> en <c>isPrime[1]</c> moeten natuurlijk <c>False </c>aangeven.)</p>
    <p>Klik hier voor een hint. Het is een uitgebreide!</p>
    <exercise label="oef-18-sieve">
      <title>Priem klaar voor takeoff</title>
      <statement>
        <p>Schrijf een programma dat een tabel definieert in <c>isPrime</c> zoals we boven beschreven (merk op dat <c>isPrime</c> een lijst is, niet een functie).   <em>De grader zal een langer dan gebruikelijke tijdlimiet toestaan om je programma uit te voeren, namelijk <alert>7 seconden</alert>. Toch zal gewoon de functie </em><c>isItPrime</c><em> uitvoeren niet snel genoeg werken.</em></p>
      </statement>
      <program interactive="activecode" language="python" label="oef-18-sieve-code" timelimit="60000">
        <code>

</code>
        <tests>
# Verborgen tests voor 18.sieve (gegenereerd door tools/harnas.py)
_OPLOSSING = 'class _Priemtabel:\n    def __len__(self):\n        return 1000001\n    def __getitem__(self, i):\n        if i &lt; 2:\n            return False\n        d = 2\n        while d * d &lt;= i:\n            if i % d == 0:\n                return False\n            d = d + 1\n        return True\nisPrime = _Priemtabel()'
_GEVALLEN = [{'aanroep': '(len(isPrime), [isPrime[i] for i in range(0, 30)], [isPrime[i] for i in (97, 100, 7919, 7917, 65537, 999983, 999999, 1000000)])'}]
_VERGELIJK = 'exact'
_TOLERANTIE = 1e-06
_HERGEBRUIK = True
_VERBODEN = []
_VERPLICHT = []
_MAXLENGTE = 0
_MAXWIJZIG = 0
_ORIGINEEL = ''

import re as _re
from unittest.gui import TestCaseGui


class _TeWeinigInvoer(Exception):
    pass


def _voer_uit(_bron, _geval, _student=False):
    _regels = [str(_r) for _r in _geval.get("invoer", [])]
    _uit = []

    def _print(*args, **kw):
        _sep = kw.get("sep", " ")
        _end = kw.get("end", "\n")
        if _sep is None:
            _sep = " "
        if _end is None:
            _end = "\n"
        _uit.append(_sep.join([str(_a) for _a in args]) + _end)

    def _input(prompt=""):
        if len(_regels) == 0:
            raise _TeWeinigInvoer("het programma vraagt meer invoer (input) dan er is")
        return _regels.pop(0)

    _g = {"__name__": "__main__", "print": _print, "input": _input}
    _fout = None
    _waarde = None
    if _student and _HERGEBRUIK:
        # Het programma van de leerling is net al uitgevoerd; gebruik die variabelen.
        try:
            _waarde = eval(_geval["aanroep"], globals())
        except Exception as _e:
            _fout = type(_e).__name__ + ": " + str(_e)
        return "", _waarde, _fout
    try:
        if _geval.get("voorcode"):
            exec(_geval["voorcode"], _g)
        exec(_bron, _g)
        if _geval.get("aanroep"):
            _waarde = eval(_geval["aanroep"], _g)
    except Exception as _e:
        _fout = type(_e).__name__ + ": " + str(_e)
    return "".join(_uit), _waarde, _fout


def _normaal(_s):
    # Spaties aan het eind van een regel tellen niet mee, en het (onzichtbare) DEL-teken ook niet.
    _r = [_x.replace(chr(127), "").rstrip() for _x in _s.split("\n")]
    while len(_r) &gt; 0 and _r[-1] == "":
        _r.pop()
    return "\n".join(_r)


def _getal(_s):
    try:
        return float(_s)
    except Exception:
        return None


def _stukken(_s):
    # Splits in getallen en overige tekens, zodat bv. "-13.333C" -&gt; ["-13.333", "C"].
    return _re.findall(r"[-+]?(?:[0-9]+[.]?[0-9]*|[.][0-9]+)(?:[eE][-+]?[0-9]+)?|[^\s0-9.+-]+|\S", _s)


def _gelijk(_a, _b):
    if _VERGELIJK == "woorden":
        return _a.split() == _b.split()
    if _VERGELIJK == "getallen":
        _wa = _stukken(_a)
        _wb = _stukken(_b)
        if len(_wa) != len(_wb):
            return False
        for _i in range(len(_wa)):
            _x = _getal(_wa[_i])
            _y = _getal(_wb[_i])
            if _x is not None and _y is not None:
                if abs(_x - _y) &gt; _TOLERANTIE * max(1.0, abs(_y)):
                    return False
            elif _wa[_i] != _wb[_i]:
                return False
        return True
    return _normaal(_a) == _normaal(_b)


def _afstand(_a, _b):
    """Bewerkingsafstand (Levenshtein) tussen twee strings."""
    _vorige = list(range(len(_b) + 1))
    for _i in range(len(_a)):
        _huidig = [_i + 1]
        for _j in range(len(_b)):
            _kosten = 0 if _a[_i] == _b[_j] else 1
            _huidig.append(min(_vorige[_j + 1] + 1, _huidig[_j] + 1, _vorige[_j] + _kosten))
        _vorige = _huidig
    return _vorige[-1]


def _zelfde(_a, _b):
    """Vergelijk returnwaarden; bij vergelijk="getallen" mogen kommagetallen iets afwijken."""
    if _VERGELIJK == "getallen" and type(_a) in (int, float) and type(_b) in (int, float) \
            and type(_a) != bool and type(_b) != bool:
        return abs(_a - _b) &lt;= _TOLERANTIE * max(1.0, abs(_b))
    if type(_a) in (list, tuple) and type(_b) == type(_a):
        if len(_a) != len(_b):
            return False
        for _i in range(len(_a)):
            if not _zelfde(_a[_i], _b[_i]):
                return False
        return True
    return type(_a) == type(_b) and _a == _b


def _html(_s):
    return str(_s).replace("&amp;", "&amp;amp;").replace("&lt;", "&amp;lt;").replace("&gt;", "&amp;gt;")


def _blok(_s):
    if _s == "":
        return "&lt;i&gt;(geen uitvoer)&lt;/i&gt;"
    return "&lt;pre style='text-align:left;margin:2px 0'&gt;" + _html(_s) + "&lt;/pre&gt;"


class _Toets(TestCaseGui):
    def test_1_programma(self):
        _bron = self.getEditorText()
        for _tekst, _melding in _VERPLICHT:
            self.assertTrue(_tekst in _bron, _melding)
        for _tekst, _melding in _VERBODEN:
            self.assertFalse(_tekst in _bron, _melding)
        if _MAXLENGTE:
            _lengte = len(_bron.rstrip())
            self.assertTrue(_lengte &lt;= _MAXLENGTE, "Je code is %d tekens lang; het mag hoogstens %d zijn."
                            % (_lengte, _MAXLENGTE))
        if _MAXWIJZIG:
            _d = _afstand(_normaal(_ORIGINEEL), _normaal(_bron))
            self.assertTrue(_d &lt;= _MAXWIJZIG, "Je hebt %d teken(s) veranderd; het mogen er hoogstens %d zijn."
                            % (_d, _MAXWIJZIG))
        _nr = 0
        for _geval in _GEVALLEN:
            _nr += 1
            _verwacht, _vwaarde, _vfout = _voer_uit(_OPLOSSING, _geval)
            _uitvoer, _waarde, _fout = _voer_uit(_bron, _geval, True)
            _info = "&lt;b&gt;Test %d&lt;/b&gt;" % _nr
            if _geval.get("voorcode"):
                _info += "&lt;br&gt;vooraf: &lt;code&gt;" + _html("; ".join(_geval["voorcode"].strip().split("\n"))) + "&lt;/code&gt;"
            if _geval.get("invoer"):
                _info += "&lt;br&gt;invoer:" + _blok("\n".join([str(_r) for _r in _geval["invoer"]]))
            if _geval.get("aanroep"):
                _info += "&lt;br&gt;aanroep: &lt;code&gt;" + _html(_geval["aanroep"]) + "&lt;/code&gt;"
            if _fout is not None and _vfout is None:
                self.appendResult(False, _fout, "", _info + "&lt;br&gt;fout: &lt;code&gt;" + _html(_fout) + "&lt;/code&gt;")
                continue
            _ok = _gelijk(_uitvoer, _verwacht)
            if _geval.get("aanroep"):
                _ok = _ok and _zelfde(_waarde, _vwaarde)
                _jij = _waarde
                _moet = _vwaarde
                if _normaal(_verwacht) != "" or _normaal(_uitvoer) != "":
                    _jij = _normaal(_uitvoer) + " / " + repr(_waarde)
                    _moet = _normaal(_verwacht) + " / " + repr(_vwaarde)
            else:
                _jij = _normaal(_uitvoer)
                _moet = _normaal(_verwacht)
            if not _ok:
                if _geval.get("aanroep"):
                    _info += "&lt;br&gt;verwachte returnwaarde: &lt;code&gt;" + _html(repr(_vwaarde)) + \
                             "&lt;/code&gt;&lt;br&gt;jouw returnwaarde: &lt;code&gt;" + _html(repr(_waarde)) + "&lt;/code&gt;"
                if not _gelijk(_uitvoer, _verwacht) or _normaal(_verwacht) != "":
                    _info += "&lt;br&gt;verwachte uitvoer:" + _blok(_normaal(_verwacht)) + \
                             "jouw uitvoer:" + _blok(_normaal(_uitvoer))
            self.appendResult(_ok, _jij, _moet, _info)


_Toets().main()

</tests>
      </program>
      <hint>
        <p>`Het is niet nodig om enige output te printen. In feite, het printen van een miljoen dingen zal tot gevolg hebben dat de tijdslimiet overschreden wordt. Je zou het eerst kunnen testen met een array van lengte 101 in plaats van 1000001 om testen gemakkelijker te maken.</p>
      </hint>
    </exercise>
    <assemblage>
      <p>Dit is de laatste oefening van de CS Circles website. Als je alle lessen hebt doorlopen: gefeliciteerd! Kijk naar de <xref ref="bijlage-bronnen" text="custom"> bronnenpagina's</xref> voor suggesties voor vervolgstappen. Veel plezier, veel succes, en programmeer ze!</p>
    </assemblage>
    <remark>
      <title>Klik hier voor een hint. Het is een uitgebreide!</title>
      <p>De oplossing die we in gedachte hebben voor dit probleem heet ook wel de "Zeef van Eratosthenes." De tabel is aanvankelijk gevuld met <c>True</c>. Vervolgens wandel je door de tabel te beginnen met index 2 en wanneer je een priemgetal tegenkomt verander je alle indices, die corresponderen met veelvouden van dat priemgetal in <c>False</c>. Lees de <url href="http://en.wikipedia.org/wiki/Sieve_of_Eratosthenes" visual="en.wikipedia.org/wiki/Sieve_of_Eratosthenes">wikipedia-pagina</url> voor alle details, en zodra je dat begrijpt kan je het omzetten in een programma in Python.</p>
    </remark>
  </section>
</chapter>
