<?xml version="1.0" encoding="UTF-8"?>
<!-- Gegenereerd door tools/convert.py uit archief/ ; niet met de hand bewerken. -->
<chapter xml:id="les-11a">
  <title>11A: Kleine letters</title>
  <introduction>
    <p><em>Les 11 heeft drie delen A, B, C die in elke willekeurige volgorde kunnen worden doorgewerkt.</em></p>
    <p>Deze les bevat een oefening waarbij je twee functies dient te schrijven: in de ene zal de andere worden gebruikt om het doel te bereiken. Het doel is om een functie <c>lowerString</c> te schrijven die alle letters van de string omzet naar <em>kleine letters</em> (Eng.: <em>lower case</em>). (A, B, C zijn <em>hoofdletters</em> (Eng.: <em>upper case</em>) en a, b, c zijn kleine letters.) Bijvoorbeeld: het resultaat van</p>
    <pre>
lowerString("This string has 9 CAPITAL letters (&amp; Punctuation)!")
</pre>
    <p>geeft</p>
    <pre>
"this string has 9 capital letters (&amp; punctuation)!"
</pre>
  </introduction>
  <section xml:id="les-11a-s1">
    <title>Stap 1: Karakters</title>
    <p>De eerste stap is om een functie <c>lowerChar(char)</c> te schrijven die in staat is de omzetting van één karakter <c>char</c> naar een kleine letter terug te geven. Ze moet het volgende doen:<ul><li><p>als het als argument meegegeven karakter <c>char</c> een hoofdletter is (tussen '<c>A</c>' en '<c>Z</c>'), dan dient ze de kleine versie van de letter terug te geven (tussen '<c>a</c>' en '<c>z</c>')</p></li><li><p>in alle andere gevallen, dient ze dezelfde <c>char</c> terug te geven die ze als argument meekreeg</p></li></ul></p>
    <p>(Om de eerste stap te kunnen uitvoeren, moet je een <c>if</c>-opdracht gebruiken, een <xref ref="les-09" text="custom"><c>and</c>-operator</xref>, en wat kennis toepassen uit <xref ref="les-07a" text="custom">de les over strings</xref>.)</p>
    <exercise label="oef-11a-lowerchar">
      <title>Kleine letters</title>
      <statement>
        <p>Schrijf een functie <c>lowerChar(char)</c> die voldoet aan bovenstaande eisen.</p>
      </statement>
      <program interactive="activecode" language="python" label="oef-11a-lowerchar-code">
        <code>

</code>
        <tests>
# Verborgen tests voor 11a.lowerchar (gegenereerd door tools/harnas.py)
_OPLOSSING = "def lowerChar(char):\n    if char &gt;= 'A' and char &lt;= 'Z':\n        return chr(ord(char) - ord('A') + ord('a'))\n    return char"
_GEVALLEN = [{'aanroep': "lowerChar('A')"}, {'aanroep': "lowerChar('Z')"}, {'aanroep': "lowerChar('M')"}, {'aanroep': "lowerChar('q')"}, {'aanroep': "lowerChar('!')"}, {'aanroep': "lowerChar('5')"}, {'aanroep': "lowerChar(' ')"}, {'aanroep': "lowerChar('[')"}, {'aanroep': "lowerChar('@')"}]
_VERGELIJK = 'exact'
_TOLERANTIE = 1e-06
_HERGEBRUIK = False
_VERBODEN = [('.lower', 'Gebruik &lt;code&gt;lower()&lt;/code&gt; niet; schrijf de omzetting zelf.')]
_VERPLICHT = []
_MAXLENGTE = 0
_MAXWIJZIG = 0
_ORIGINEEL = ''

import re as _re
from unittest.gui import TestCaseGui


class _TeWeinigInvoer(Exception):
    pass


def _voer_uit(_bron, _geval, _student=False):
    _regels = [str(_r) for _r in _geval.get("invoer", [])]
    _uit = []

    def _print(*args, **kw):
        _sep = kw.get("sep", " ")
        _end = kw.get("end", "\n")
        if _sep is None:
            _sep = " "
        if _end is None:
            _end = "\n"
        _uit.append(_sep.join([str(_a) for _a in args]) + _end)

    def _input(prompt=""):
        if len(_regels) == 0:
            raise _TeWeinigInvoer("het programma vraagt meer invoer (input) dan er is")
        return _regels.pop(0)

    _g = {"__name__": "__main__", "print": _print, "input": _input}
    _fout = None
    _waarde = None
    if _student and _HERGEBRUIK:
        # Het programma van de leerling is net al uitgevoerd; gebruik die variabelen.
        try:
            _waarde = eval(_geval["aanroep"], globals())
        except Exception as _e:
            _fout = type(_e).__name__ + ": " + str(_e)
        return "", _waarde, _fout
    try:
        if _geval.get("voorcode"):
            exec(_geval["voorcode"], _g)
        exec(_bron, _g)
        if _geval.get("aanroep"):
            _waarde = eval(_geval["aanroep"], _g)
    except Exception as _e:
        _fout = type(_e).__name__ + ": " + str(_e)
    return "".join(_uit), _waarde, _fout


def _normaal(_s):
    # Spaties aan het eind van een regel tellen niet mee, en het (onzichtbare) DEL-teken ook niet.
    _r = [_x.replace(chr(127), "").rstrip() for _x in _s.split("\n")]
    while len(_r) &gt; 0 and _r[-1] == "":
        _r.pop()
    return "\n".join(_r)


def _getal(_s):
    try:
        return float(_s)
    except Exception:
        return None


def _stukken(_s):
    # Splits in getallen en overige tekens, zodat bv. "-13.333C" -&gt; ["-13.333", "C"].
    return _re.findall(r"[-+]?(?:[0-9]+[.]?[0-9]*|[.][0-9]+)(?:[eE][-+]?[0-9]+)?|[^\s0-9.+-]+|\S", _s)


def _gelijk(_a, _b):
    if _VERGELIJK == "woorden":
        return _a.split() == _b.split()
    if _VERGELIJK == "getallen":
        _wa = _stukken(_a)
        _wb = _stukken(_b)
        if len(_wa) != len(_wb):
            return False
        for _i in range(len(_wa)):
            _x = _getal(_wa[_i])
            _y = _getal(_wb[_i])
            if _x is not None and _y is not None:
                if abs(_x - _y) &gt; _TOLERANTIE * max(1.0, abs(_y)):
                    return False
            elif _wa[_i] != _wb[_i]:
                return False
        return True
    return _normaal(_a) == _normaal(_b)


def _afstand(_a, _b):
    """Bewerkingsafstand (Levenshtein) tussen twee strings."""
    _vorige = list(range(len(_b) + 1))
    for _i in range(len(_a)):
        _huidig = [_i + 1]
        for _j in range(len(_b)):
            _kosten = 0 if _a[_i] == _b[_j] else 1
            _huidig.append(min(_vorige[_j + 1] + 1, _huidig[_j] + 1, _vorige[_j] + _kosten))
        _vorige = _huidig
    return _vorige[-1]


def _zelfde(_a, _b):
    """Vergelijk returnwaarden; bij vergelijk="getallen" mogen kommagetallen iets afwijken."""
    if _VERGELIJK == "getallen" and type(_a) in (int, float) and type(_b) in (int, float) \
            and type(_a) != bool and type(_b) != bool:
        return abs(_a - _b) &lt;= _TOLERANTIE * max(1.0, abs(_b))
    if type(_a) in (list, tuple) and type(_b) == type(_a):
        if len(_a) != len(_b):
            return False
        for _i in range(len(_a)):
            if not _zelfde(_a[_i], _b[_i]):
                return False
        return True
    return type(_a) == type(_b) and _a == _b


def _html(_s):
    return str(_s).replace("&amp;", "&amp;amp;").replace("&lt;", "&amp;lt;").replace("&gt;", "&amp;gt;")


def _blok(_s):
    if _s == "":
        return "&lt;i&gt;(geen uitvoer)&lt;/i&gt;"
    return "&lt;pre style='text-align:left;margin:2px 0'&gt;" + _html(_s) + "&lt;/pre&gt;"


class _Toets(TestCaseGui):
    def test_1_programma(self):
        _bron = self.getEditorText()
        for _tekst, _melding in _VERPLICHT:
            self.assertTrue(_tekst in _bron, _melding)
        for _tekst, _melding in _VERBODEN:
            self.assertFalse(_tekst in _bron, _melding)
        if _MAXLENGTE:
            _lengte = len(_bron.rstrip())
            self.assertTrue(_lengte &lt;= _MAXLENGTE, "Je code is %d tekens lang; het mag hoogstens %d zijn."
                            % (_lengte, _MAXLENGTE))
        if _MAXWIJZIG:
            _d = _afstand(_normaal(_ORIGINEEL), _normaal(_bron))
            self.assertTrue(_d &lt;= _MAXWIJZIG, "Je hebt %d teken(s) veranderd; het mogen er hoogstens %d zijn."
                            % (_d, _MAXWIJZIG))
        _nr = 0
        for _geval in _GEVALLEN:
            _nr += 1
            _verwacht, _vwaarde, _vfout = _voer_uit(_OPLOSSING, _geval)
            _uitvoer, _waarde, _fout = _voer_uit(_bron, _geval, True)
            _info = "&lt;b&gt;Test %d&lt;/b&gt;" % _nr
            if _geval.get("voorcode"):
                _info += "&lt;br&gt;vooraf: &lt;code&gt;" + _html("; ".join(_geval["voorcode"].strip().split("\n"))) + "&lt;/code&gt;"
            if _geval.get("invoer"):
                _info += "&lt;br&gt;invoer:" + _blok("\n".join([str(_r) for _r in _geval["invoer"]]))
            if _geval.get("aanroep"):
                _info += "&lt;br&gt;aanroep: &lt;code&gt;" + _html(_geval["aanroep"]) + "&lt;/code&gt;"
            if _fout is not None and _vfout is None:
                self.appendResult(False, _fout, "", _info + "&lt;br&gt;fout: &lt;code&gt;" + _html(_fout) + "&lt;/code&gt;")
                continue
            _ok = _gelijk(_uitvoer, _verwacht)
            if _geval.get("aanroep"):
                _ok = _ok and _zelfde(_waarde, _vwaarde)
                _jij = _waarde
                _moet = _vwaarde
                if _normaal(_verwacht) != "" or _normaal(_uitvoer) != "":
                    _jij = _normaal(_uitvoer) + " / " + repr(_waarde)
                    _moet = _normaal(_verwacht) + " / " + repr(_vwaarde)
            else:
                _jij = _normaal(_uitvoer)
                _moet = _normaal(_verwacht)
            if not _ok:
                if _geval.get("aanroep"):
                    _info += "&lt;br&gt;verwachte returnwaarde: &lt;code&gt;" + _html(repr(_vwaarde)) + \
                             "&lt;/code&gt;&lt;br&gt;jouw returnwaarde: &lt;code&gt;" + _html(repr(_waarde)) + "&lt;/code&gt;"
                if not _gelijk(_uitvoer, _verwacht) or _normaal(_verwacht) != "":
                    _info += "&lt;br&gt;verwachte uitvoer:" + _blok(_normaal(_verwacht)) + \
                             "jouw uitvoer:" + _blok(_normaal(_uitvoer))
            self.appendResult(_ok, _jij, _moet, _info)


_Toets().main()

</tests>
      </program>
    </exercise>
  </section>
  <section xml:id="les-11a-s2">
    <title>Stap 2: Strings</title>
    <p>Nu moet je een tweede functie schrijven genaamd <c>lowerString(string)</c> die als resultaat de hele string in kleine letters teruggeeft, door de functie <c>lowerChar</c> voor elk karakter aan te roepen. We stellen voor dat als volgt te doen:<ul><li><p>begin met het kopiëren van de definitie van <c>lowerChar(char)</c> uit het eerste deel</p></li><li><p>schrijf dan een tweede functie, <c>lowerString(string)</c></p><p><ul><li><p>op de eerste regel van lowerString, definieer een variabele <c>result = ""</c> en maak deze gelijk aan de lege string</p></li><li><p>gebruik een <xref ref="les-07c" text="custom">for-lus</xref> met <c>i</c> en neem op <c>result = result + lowerChar(string[i])</c></p></li><li><p>zorg er tenslotte voor dat <c>result</c> wordt teruggegeven</p></li></ul></p></li></ul></p>
    <exercise label="oef-11a-lowerstring">
      <title>Lower-case Strings</title>
      <statement>
        <p>Schrijf een functie <c>lowerString(string)</c> die als resultaat de omzetting van de <c>string</c> in kleine letters teruggeeft.</p>
      </statement>
      <program interactive="activecode" language="python" label="oef-11a-lowerstring-code">
        <code>
# first, copy your definition of lowerChar() here
# then define lowerString(string)
</code>
        <tests>
# Verborgen tests voor 11a.lowerstring (gegenereerd door tools/harnas.py)
_OPLOSSING = 'def lowerChar(char):\n    if char &gt;= \'A\' and char &lt;= \'Z\':\n        return chr(ord(char) - ord(\'A\') + ord(\'a\'))\n    return char\n\ndef lowerString(string):\n    result = ""\n    for i in range(0, len(string)):\n        result = result + lowerChar(string[i])\n    return result'
_GEVALLEN = [{'aanroep': "lowerString('Hallo Wereld!')"}, {'aanroep': "lowerString('')"}, {'aanroep': "lowerString('ABC xyz 123')"}, {'aanroep': "lowerString('Python@CS-Circles')"}]
_VERGELIJK = 'exact'
_TOLERANTIE = 1e-06
_HERGEBRUIK = False
_VERBODEN = [('.lower', 'Gebruik &lt;code&gt;lower()&lt;/code&gt; niet; gebruik je eigen functie &lt;code&gt;lowerChar&lt;/code&gt;.')]
_VERPLICHT = []
_MAXLENGTE = 0
_MAXWIJZIG = 0
_ORIGINEEL = ''

import re as _re
from unittest.gui import TestCaseGui


class _TeWeinigInvoer(Exception):
    pass


def _voer_uit(_bron, _geval, _student=False):
    _regels = [str(_r) for _r in _geval.get("invoer", [])]
    _uit = []

    def _print(*args, **kw):
        _sep = kw.get("sep", " ")
        _end = kw.get("end", "\n")
        if _sep is None:
            _sep = " "
        if _end is None:
            _end = "\n"
        _uit.append(_sep.join([str(_a) for _a in args]) + _end)

    def _input(prompt=""):
        if len(_regels) == 0:
            raise _TeWeinigInvoer("het programma vraagt meer invoer (input) dan er is")
        return _regels.pop(0)

    _g = {"__name__": "__main__", "print": _print, "input": _input}
    _fout = None
    _waarde = None
    if _student and _HERGEBRUIK:
        # Het programma van de leerling is net al uitgevoerd; gebruik die variabelen.
        try:
            _waarde = eval(_geval["aanroep"], globals())
        except Exception as _e:
            _fout = type(_e).__name__ + ": " + str(_e)
        return "", _waarde, _fout
    try:
        if _geval.get("voorcode"):
            exec(_geval["voorcode"], _g)
        exec(_bron, _g)
        if _geval.get("aanroep"):
            _waarde = eval(_geval["aanroep"], _g)
    except Exception as _e:
        _fout = type(_e).__name__ + ": " + str(_e)
    return "".join(_uit), _waarde, _fout


def _normaal(_s):
    # Spaties aan het eind van een regel tellen niet mee, en het (onzichtbare) DEL-teken ook niet.
    _r = [_x.replace(chr(127), "").rstrip() for _x in _s.split("\n")]
    while len(_r) &gt; 0 and _r[-1] == "":
        _r.pop()
    return "\n".join(_r)


def _getal(_s):
    try:
        return float(_s)
    except Exception:
        return None


def _stukken(_s):
    # Splits in getallen en overige tekens, zodat bv. "-13.333C" -&gt; ["-13.333", "C"].
    return _re.findall(r"[-+]?(?:[0-9]+[.]?[0-9]*|[.][0-9]+)(?:[eE][-+]?[0-9]+)?|[^\s0-9.+-]+|\S", _s)


def _gelijk(_a, _b):
    if _VERGELIJK == "woorden":
        return _a.split() == _b.split()
    if _VERGELIJK == "getallen":
        _wa = _stukken(_a)
        _wb = _stukken(_b)
        if len(_wa) != len(_wb):
            return False
        for _i in range(len(_wa)):
            _x = _getal(_wa[_i])
            _y = _getal(_wb[_i])
            if _x is not None and _y is not None:
                if abs(_x - _y) &gt; _TOLERANTIE * max(1.0, abs(_y)):
                    return False
            elif _wa[_i] != _wb[_i]:
                return False
        return True
    return _normaal(_a) == _normaal(_b)


def _afstand(_a, _b):
    """Bewerkingsafstand (Levenshtein) tussen twee strings."""
    _vorige = list(range(len(_b) + 1))
    for _i in range(len(_a)):
        _huidig = [_i + 1]
        for _j in range(len(_b)):
            _kosten = 0 if _a[_i] == _b[_j] else 1
            _huidig.append(min(_vorige[_j + 1] + 1, _huidig[_j] + 1, _vorige[_j] + _kosten))
        _vorige = _huidig
    return _vorige[-1]


def _zelfde(_a, _b):
    """Vergelijk returnwaarden; bij vergelijk="getallen" mogen kommagetallen iets afwijken."""
    if _VERGELIJK == "getallen" and type(_a) in (int, float) and type(_b) in (int, float) \
            and type(_a) != bool and type(_b) != bool:
        return abs(_a - _b) &lt;= _TOLERANTIE * max(1.0, abs(_b))
    if type(_a) in (list, tuple) and type(_b) == type(_a):
        if len(_a) != len(_b):
            return False
        for _i in range(len(_a)):
            if not _zelfde(_a[_i], _b[_i]):
                return False
        return True
    return type(_a) == type(_b) and _a == _b


def _html(_s):
    return str(_s).replace("&amp;", "&amp;amp;").replace("&lt;", "&amp;lt;").replace("&gt;", "&amp;gt;")


def _blok(_s):
    if _s == "":
        return "&lt;i&gt;(geen uitvoer)&lt;/i&gt;"
    return "&lt;pre style='text-align:left;margin:2px 0'&gt;" + _html(_s) + "&lt;/pre&gt;"


class _Toets(TestCaseGui):
    def test_1_programma(self):
        _bron = self.getEditorText()
        for _tekst, _melding in _VERPLICHT:
            self.assertTrue(_tekst in _bron, _melding)
        for _tekst, _melding in _VERBODEN:
            self.assertFalse(_tekst in _bron, _melding)
        if _MAXLENGTE:
            _lengte = len(_bron.rstrip())
            self.assertTrue(_lengte &lt;= _MAXLENGTE, "Je code is %d tekens lang; het mag hoogstens %d zijn."
                            % (_lengte, _MAXLENGTE))
        if _MAXWIJZIG:
            _d = _afstand(_normaal(_ORIGINEEL), _normaal(_bron))
            self.assertTrue(_d &lt;= _MAXWIJZIG, "Je hebt %d teken(s) veranderd; het mogen er hoogstens %d zijn."
                            % (_d, _MAXWIJZIG))
        _nr = 0
        for _geval in _GEVALLEN:
            _nr += 1
            _verwacht, _vwaarde, _vfout = _voer_uit(_OPLOSSING, _geval)
            _uitvoer, _waarde, _fout = _voer_uit(_bron, _geval, True)
            _info = "&lt;b&gt;Test %d&lt;/b&gt;" % _nr
            if _geval.get("voorcode"):
                _info += "&lt;br&gt;vooraf: &lt;code&gt;" + _html("; ".join(_geval["voorcode"].strip().split("\n"))) + "&lt;/code&gt;"
            if _geval.get("invoer"):
                _info += "&lt;br&gt;invoer:" + _blok("\n".join([str(_r) for _r in _geval["invoer"]]))
            if _geval.get("aanroep"):
                _info += "&lt;br&gt;aanroep: &lt;code&gt;" + _html(_geval["aanroep"]) + "&lt;/code&gt;"
            if _fout is not None and _vfout is None:
                self.appendResult(False, _fout, "", _info + "&lt;br&gt;fout: &lt;code&gt;" + _html(_fout) + "&lt;/code&gt;")
                continue
            _ok = _gelijk(_uitvoer, _verwacht)
            if _geval.get("aanroep"):
                _ok = _ok and _zelfde(_waarde, _vwaarde)
                _jij = _waarde
                _moet = _vwaarde
                if _normaal(_verwacht) != "" or _normaal(_uitvoer) != "":
                    _jij = _normaal(_uitvoer) + " / " + repr(_waarde)
                    _moet = _normaal(_verwacht) + " / " + repr(_vwaarde)
            else:
                _jij = _normaal(_uitvoer)
                _moet = _normaal(_verwacht)
            if not _ok:
                if _geval.get("aanroep"):
                    _info += "&lt;br&gt;verwachte returnwaarde: &lt;code&gt;" + _html(repr(_vwaarde)) + \
                             "&lt;/code&gt;&lt;br&gt;jouw returnwaarde: &lt;code&gt;" + _html(repr(_waarde)) + "&lt;/code&gt;"
                if not _gelijk(_uitvoer, _verwacht) or _normaal(_verwacht) != "":
                    _info += "&lt;br&gt;verwachte uitvoer:" + _blok(_normaal(_verwacht)) + \
                             "jouw uitvoer:" + _blok(_normaal(_uitvoer))
            self.appendResult(_ok, _jij, _moet, _info)


_Toets().main()

</tests>
      </program>
    </exercise>
    <assemblage>
      <p>Later zul je <c>string.lower()</c> leren kennen, wat een ingebouwd element is van Python en een string omzet naar een string met alleen kleine letters.</p>
    </assemblage>
  </section>
</chapter>
