<?xml version="1.0" encoding="UTF-8"?>
<!-- Gegenereerd door tools/convert.py uit archief/ ; niet met de hand bewerken. -->
<chapter xml:id="les-11c">
  <title>11C: Meetkunde</title>
  <introduction>
    <p><em>Les 11 kent drie delen A, B, C die in een willekeurige volgorde kunnen worden doorgewerkt.</em></p>
    <p>In deze oefening zullen we vier functies realiseren die bij berekeningen in de meetkunde van pas komen:<ul><li><p>een functie die de lengte uitrekent van de schuine zijde van een rechthoekige driehoek,</p></li><li><p>een functie die de omtrek uitrekent van een rechthoekige driehoek,</p></li><li><p>een functie die de afstand tussen twee punten in het platte vlak uitrekent,</p></li><li><p>een functie die de omtrek uitrekent van een willekeurige driehoek.</p></li></ul></p>
  </introduction>
  <section xml:id="les-11c-s1">
    <title>Hypotenusa (schuine zijde)</title>
    <p>Je ziet een <em>rechthoekige driehoekig</em> in het plaatje rechts. Volgens de definitie is één van de hoeken van de rechthoekige driehoek gelijk aan 90 graden (een <em>rechte hoek</em>). De driehoek heeft drie zijden. De zijde tegenover de rechte hoek heeft een speciale naam: de <em>hypotenusa.</em></p>
    <image source="images/11c-hypotenuse.png" width="30%">
      <shortdescription>hypotenuse</shortdescription>
    </image>
    <p>Zoals in de tekening te zien is laten we <em>a</em> en <em>b</em> de lengte zijn van de rechthoekszijden, en <em>c</em> de lengte van de hypotenusa. De beroemde Stelling van Pythagoras zegt nu dat:</p>
    <p><m>\displaystyle{a^2+b^2=c^2}</m></p>
    <p>In de eerste opgave moet je deze stelling in een functie omzetten die, wanneer <em>a</em> en <em>b</em> gegeven zijn, de lengte van de hypotenusa berekent.</p>
    <exercise label="oef-11c-hypotenuse">
      <title>Hypotenusa</title>
      <statement>
        <p>Definieer een functie <c>hypotenuse(a, b)</c> die de lengte van de hypotenusa <em>c</em> teruggeeft, wanneer de andere twee zijden lengte <em>a</em> en <em>b</em> hebben.</p>
      </statement>
      <program interactive="activecode" language="python" label="oef-11c-hypotenuse-code">
        <code>

</code>
        <tests>
# Verborgen tests voor 11c.hypotenuse (gegenereerd door tools/harnas.py)
_OPLOSSING = 'import math\ndef hypotenuse(a, b):\n    return math.sqrt(a * a + b * b)'
_GEVALLEN = [{'aanroep': 'hypotenuse(3, 4)'}, {'aanroep': 'hypotenuse(5, 12)'}, {'aanroep': 'hypotenuse(1, 1)'}, {'aanroep': 'hypotenuse(0, 7)'}, {'aanroep': 'hypotenuse(2.5, 6)'}]
_VERGELIJK = 'getallen'
_TOLERANTIE = 1e-06
_HERGEBRUIK = False
_VERBODEN = []
_VERPLICHT = []
_MAXLENGTE = 0
_MAXWIJZIG = 0
_ORIGINEEL = ''

import re as _re
from unittest.gui import TestCaseGui


class _TeWeinigInvoer(Exception):
    pass


def _voer_uit(_bron, _geval, _student=False):
    _regels = [str(_r) for _r in _geval.get("invoer", [])]
    _uit = []

    def _print(*args, **kw):
        _sep = kw.get("sep", " ")
        _end = kw.get("end", "\n")
        if _sep is None:
            _sep = " "
        if _end is None:
            _end = "\n"
        _uit.append(_sep.join([str(_a) for _a in args]) + _end)

    def _input(prompt=""):
        if len(_regels) == 0:
            raise _TeWeinigInvoer("het programma vraagt meer invoer (input) dan er is")
        return _regels.pop(0)

    _g = {"__name__": "__main__", "print": _print, "input": _input}
    _fout = None
    _waarde = None
    if _student and _HERGEBRUIK:
        # Het programma van de leerling is net al uitgevoerd; gebruik die variabelen.
        try:
            _waarde = eval(_geval["aanroep"], globals())
        except Exception as _e:
            _fout = type(_e).__name__ + ": " + str(_e)
        return "", _waarde, _fout
    try:
        if _geval.get("voorcode"):
            exec(_geval["voorcode"], _g)
        exec(_bron, _g)
        if _geval.get("aanroep"):
            _waarde = eval(_geval["aanroep"], _g)
    except Exception as _e:
        _fout = type(_e).__name__ + ": " + str(_e)
    return "".join(_uit), _waarde, _fout


def _normaal(_s):
    # Spaties aan het eind van een regel tellen niet mee, en het (onzichtbare) DEL-teken ook niet.
    _r = [_x.replace(chr(127), "").rstrip() for _x in _s.split("\n")]
    while len(_r) &gt; 0 and _r[-1] == "":
        _r.pop()
    return "\n".join(_r)


def _getal(_s):
    try:
        return float(_s)
    except Exception:
        return None


def _stukken(_s):
    # Splits in getallen en overige tekens, zodat bv. "-13.333C" -&gt; ["-13.333", "C"].
    return _re.findall(r"[-+]?(?:[0-9]+[.]?[0-9]*|[.][0-9]+)(?:[eE][-+]?[0-9]+)?|[^\s0-9.+-]+|\S", _s)


def _gelijk(_a, _b):
    if _VERGELIJK == "woorden":
        return _a.split() == _b.split()
    if _VERGELIJK == "getallen":
        _wa = _stukken(_a)
        _wb = _stukken(_b)
        if len(_wa) != len(_wb):
            return False
        for _i in range(len(_wa)):
            _x = _getal(_wa[_i])
            _y = _getal(_wb[_i])
            if _x is not None and _y is not None:
                if abs(_x - _y) &gt; _TOLERANTIE * max(1.0, abs(_y)):
                    return False
            elif _wa[_i] != _wb[_i]:
                return False
        return True
    return _normaal(_a) == _normaal(_b)


def _afstand(_a, _b):
    """Bewerkingsafstand (Levenshtein) tussen twee strings."""
    _vorige = list(range(len(_b) + 1))
    for _i in range(len(_a)):
        _huidig = [_i + 1]
        for _j in range(len(_b)):
            _kosten = 0 if _a[_i] == _b[_j] else 1
            _huidig.append(min(_vorige[_j + 1] + 1, _huidig[_j] + 1, _vorige[_j] + _kosten))
        _vorige = _huidig
    return _vorige[-1]


def _zelfde(_a, _b):
    """Vergelijk returnwaarden; bij vergelijk="getallen" mogen kommagetallen iets afwijken."""
    if _VERGELIJK == "getallen" and type(_a) in (int, float) and type(_b) in (int, float) \
            and type(_a) != bool and type(_b) != bool:
        return abs(_a - _b) &lt;= _TOLERANTIE * max(1.0, abs(_b))
    if type(_a) in (list, tuple) and type(_b) == type(_a):
        if len(_a) != len(_b):
            return False
        for _i in range(len(_a)):
            if not _zelfde(_a[_i], _b[_i]):
                return False
        return True
    return type(_a) == type(_b) and _a == _b


def _html(_s):
    return str(_s).replace("&amp;", "&amp;amp;").replace("&lt;", "&amp;lt;").replace("&gt;", "&amp;gt;")


def _blok(_s):
    if _s == "":
        return "&lt;i&gt;(geen uitvoer)&lt;/i&gt;"
    return "&lt;pre style='text-align:left;margin:2px 0'&gt;" + _html(_s) + "&lt;/pre&gt;"


class _Toets(TestCaseGui):
    def test_1_programma(self):
        _bron = self.getEditorText()
        for _tekst, _melding in _VERPLICHT:
            self.assertTrue(_tekst in _bron, _melding)
        for _tekst, _melding in _VERBODEN:
            self.assertFalse(_tekst in _bron, _melding)
        if _MAXLENGTE:
            _lengte = len(_bron.rstrip())
            self.assertTrue(_lengte &lt;= _MAXLENGTE, "Je code is %d tekens lang; het mag hoogstens %d zijn."
                            % (_lengte, _MAXLENGTE))
        if _MAXWIJZIG:
            _d = _afstand(_normaal(_ORIGINEEL), _normaal(_bron))
            self.assertTrue(_d &lt;= _MAXWIJZIG, "Je hebt %d teken(s) veranderd; het mogen er hoogstens %d zijn."
                            % (_d, _MAXWIJZIG))
        _nr = 0
        for _geval in _GEVALLEN:
            _nr += 1
            _verwacht, _vwaarde, _vfout = _voer_uit(_OPLOSSING, _geval)
            _uitvoer, _waarde, _fout = _voer_uit(_bron, _geval, True)
            _info = "&lt;b&gt;Test %d&lt;/b&gt;" % _nr
            if _geval.get("voorcode"):
                _info += "&lt;br&gt;vooraf: &lt;code&gt;" + _html("; ".join(_geval["voorcode"].strip().split("\n"))) + "&lt;/code&gt;"
            if _geval.get("invoer"):
                _info += "&lt;br&gt;invoer:" + _blok("\n".join([str(_r) for _r in _geval["invoer"]]))
            if _geval.get("aanroep"):
                _info += "&lt;br&gt;aanroep: &lt;code&gt;" + _html(_geval["aanroep"]) + "&lt;/code&gt;"
            if _fout is not None and _vfout is None:
                self.appendResult(False, _fout, "", _info + "&lt;br&gt;fout: &lt;code&gt;" + _html(_fout) + "&lt;/code&gt;")
                continue
            _ok = _gelijk(_uitvoer, _verwacht)
            if _geval.get("aanroep"):
                _ok = _ok and _zelfde(_waarde, _vwaarde)
                _jij = _waarde
                _moet = _vwaarde
                if _normaal(_verwacht) != "" or _normaal(_uitvoer) != "":
                    _jij = _normaal(_uitvoer) + " / " + repr(_waarde)
                    _moet = _normaal(_verwacht) + " / " + repr(_vwaarde)
            else:
                _jij = _normaal(_uitvoer)
                _moet = _normaal(_verwacht)
            if not _ok:
                if _geval.get("aanroep"):
                    _info += "&lt;br&gt;verwachte returnwaarde: &lt;code&gt;" + _html(repr(_vwaarde)) + \
                             "&lt;/code&gt;&lt;br&gt;jouw returnwaarde: &lt;code&gt;" + _html(repr(_waarde)) + "&lt;/code&gt;"
                if not _gelijk(_uitvoer, _verwacht) or _normaal(_verwacht) != "":
                    _info += "&lt;br&gt;verwachte uitvoer:" + _blok(_normaal(_verwacht)) + \
                             "jouw uitvoer:" + _blok(_normaal(_uitvoer))
            self.appendResult(_ok, _jij, _moet, _info)


_Toets().main()

</tests>
      </program>
      <hint>
        <p>Je zult <c>import math</c> nodig hebben en <c>math.sqrt()</c> op iets moeten loslaten.</p>
      </hint>
    </exercise>
  </section>
  <section xml:id="les-11c-s2">
    <title>Omtrek</title>
    <p>Denk eraan dat de <em>omtrek</em> van een driehoek de som van de lengtes van de zijden is. Dus in bovenstaande tekening is de omtrek <em>a+b+c</em>. <alert> In je programma mag je aannemen</alert> dat een correcte versie van <c>hypotenusa</c> reeds is gedefinieerd (het is niet nodig om het vanuit de eerste box naar de tweede te kopiëren).</p>
    <exercise label="oef-11c-rightperimeter">
      <title>De driehoeken zijn rechthoekig</title>
      <statement>
        <p>Definieer een functie <c>rightTrianglePerimeter(a, b)</c> die de omtrek teruggeeft van een rechthoekige driehoek met rechthoekszijden a en b. Maak hierbij gebruik van <c>hypotenuse(a, b)</c>.</p>
      </statement>
      <program interactive="activecode" language="python" label="oef-11c-rightperimeter-code">
        <code>

</code>
        <tests>
# Verborgen tests voor 11c.rightperimeter (gegenereerd door tools/harnas.py)
_OPLOSSING = 'import math\ndef hypotenuse(a, b):\n    return math.sqrt(a * a + b * b)\n\ndef rightTrianglePerimeter(a, b):\n    return a + b + hypotenuse(a, b)'
_GEVALLEN = [{'aanroep': 'rightTrianglePerimeter(3, 4)'}, {'aanroep': 'rightTrianglePerimeter(5, 12)'}, {'aanroep': 'rightTrianglePerimeter(1, 1)'}, {'aanroep': 'rightTrianglePerimeter(8, 15)'}]
_VERGELIJK = 'getallen'
_TOLERANTIE = 1e-06
_HERGEBRUIK = False
_VERBODEN = []
_VERPLICHT = []
_MAXLENGTE = 0
_MAXWIJZIG = 0
_ORIGINEEL = ''

import re as _re
from unittest.gui import TestCaseGui


class _TeWeinigInvoer(Exception):
    pass


def _voer_uit(_bron, _geval, _student=False):
    _regels = [str(_r) for _r in _geval.get("invoer", [])]
    _uit = []

    def _print(*args, **kw):
        _sep = kw.get("sep", " ")
        _end = kw.get("end", "\n")
        if _sep is None:
            _sep = " "
        if _end is None:
            _end = "\n"
        _uit.append(_sep.join([str(_a) for _a in args]) + _end)

    def _input(prompt=""):
        if len(_regels) == 0:
            raise _TeWeinigInvoer("het programma vraagt meer invoer (input) dan er is")
        return _regels.pop(0)

    _g = {"__name__": "__main__", "print": _print, "input": _input}
    _fout = None
    _waarde = None
    if _student and _HERGEBRUIK:
        # Het programma van de leerling is net al uitgevoerd; gebruik die variabelen.
        try:
            _waarde = eval(_geval["aanroep"], globals())
        except Exception as _e:
            _fout = type(_e).__name__ + ": " + str(_e)
        return "", _waarde, _fout
    try:
        if _geval.get("voorcode"):
            exec(_geval["voorcode"], _g)
        exec(_bron, _g)
        if _geval.get("aanroep"):
            _waarde = eval(_geval["aanroep"], _g)
    except Exception as _e:
        _fout = type(_e).__name__ + ": " + str(_e)
    return "".join(_uit), _waarde, _fout


def _normaal(_s):
    # Spaties aan het eind van een regel tellen niet mee, en het (onzichtbare) DEL-teken ook niet.
    _r = [_x.replace(chr(127), "").rstrip() for _x in _s.split("\n")]
    while len(_r) &gt; 0 and _r[-1] == "":
        _r.pop()
    return "\n".join(_r)


def _getal(_s):
    try:
        return float(_s)
    except Exception:
        return None


def _stukken(_s):
    # Splits in getallen en overige tekens, zodat bv. "-13.333C" -&gt; ["-13.333", "C"].
    return _re.findall(r"[-+]?(?:[0-9]+[.]?[0-9]*|[.][0-9]+)(?:[eE][-+]?[0-9]+)?|[^\s0-9.+-]+|\S", _s)


def _gelijk(_a, _b):
    if _VERGELIJK == "woorden":
        return _a.split() == _b.split()
    if _VERGELIJK == "getallen":
        _wa = _stukken(_a)
        _wb = _stukken(_b)
        if len(_wa) != len(_wb):
            return False
        for _i in range(len(_wa)):
            _x = _getal(_wa[_i])
            _y = _getal(_wb[_i])
            if _x is not None and _y is not None:
                if abs(_x - _y) &gt; _TOLERANTIE * max(1.0, abs(_y)):
                    return False
            elif _wa[_i] != _wb[_i]:
                return False
        return True
    return _normaal(_a) == _normaal(_b)


def _afstand(_a, _b):
    """Bewerkingsafstand (Levenshtein) tussen twee strings."""
    _vorige = list(range(len(_b) + 1))
    for _i in range(len(_a)):
        _huidig = [_i + 1]
        for _j in range(len(_b)):
            _kosten = 0 if _a[_i] == _b[_j] else 1
            _huidig.append(min(_vorige[_j + 1] + 1, _huidig[_j] + 1, _vorige[_j] + _kosten))
        _vorige = _huidig
    return _vorige[-1]


def _zelfde(_a, _b):
    """Vergelijk returnwaarden; bij vergelijk="getallen" mogen kommagetallen iets afwijken."""
    if _VERGELIJK == "getallen" and type(_a) in (int, float) and type(_b) in (int, float) \
            and type(_a) != bool and type(_b) != bool:
        return abs(_a - _b) &lt;= _TOLERANTIE * max(1.0, abs(_b))
    if type(_a) in (list, tuple) and type(_b) == type(_a):
        if len(_a) != len(_b):
            return False
        for _i in range(len(_a)):
            if not _zelfde(_a[_i], _b[_i]):
                return False
        return True
    return type(_a) == type(_b) and _a == _b


def _html(_s):
    return str(_s).replace("&amp;", "&amp;amp;").replace("&lt;", "&amp;lt;").replace("&gt;", "&amp;gt;")


def _blok(_s):
    if _s == "":
        return "&lt;i&gt;(geen uitvoer)&lt;/i&gt;"
    return "&lt;pre style='text-align:left;margin:2px 0'&gt;" + _html(_s) + "&lt;/pre&gt;"


class _Toets(TestCaseGui):
    def test_1_programma(self):
        _bron = self.getEditorText()
        for _tekst, _melding in _VERPLICHT:
            self.assertTrue(_tekst in _bron, _melding)
        for _tekst, _melding in _VERBODEN:
            self.assertFalse(_tekst in _bron, _melding)
        if _MAXLENGTE:
            _lengte = len(_bron.rstrip())
            self.assertTrue(_lengte &lt;= _MAXLENGTE, "Je code is %d tekens lang; het mag hoogstens %d zijn."
                            % (_lengte, _MAXLENGTE))
        if _MAXWIJZIG:
            _d = _afstand(_normaal(_ORIGINEEL), _normaal(_bron))
            self.assertTrue(_d &lt;= _MAXWIJZIG, "Je hebt %d teken(s) veranderd; het mogen er hoogstens %d zijn."
                            % (_d, _MAXWIJZIG))
        _nr = 0
        for _geval in _GEVALLEN:
            _nr += 1
            _verwacht, _vwaarde, _vfout = _voer_uit(_OPLOSSING, _geval)
            _uitvoer, _waarde, _fout = _voer_uit(_bron, _geval, True)
            _info = "&lt;b&gt;Test %d&lt;/b&gt;" % _nr
            if _geval.get("voorcode"):
                _info += "&lt;br&gt;vooraf: &lt;code&gt;" + _html("; ".join(_geval["voorcode"].strip().split("\n"))) + "&lt;/code&gt;"
            if _geval.get("invoer"):
                _info += "&lt;br&gt;invoer:" + _blok("\n".join([str(_r) for _r in _geval["invoer"]]))
            if _geval.get("aanroep"):
                _info += "&lt;br&gt;aanroep: &lt;code&gt;" + _html(_geval["aanroep"]) + "&lt;/code&gt;"
            if _fout is not None and _vfout is None:
                self.appendResult(False, _fout, "", _info + "&lt;br&gt;fout: &lt;code&gt;" + _html(_fout) + "&lt;/code&gt;")
                continue
            _ok = _gelijk(_uitvoer, _verwacht)
            if _geval.get("aanroep"):
                _ok = _ok and _zelfde(_waarde, _vwaarde)
                _jij = _waarde
                _moet = _vwaarde
                if _normaal(_verwacht) != "" or _normaal(_uitvoer) != "":
                    _jij = _normaal(_uitvoer) + " / " + repr(_waarde)
                    _moet = _normaal(_verwacht) + " / " + repr(_vwaarde)
            else:
                _jij = _normaal(_uitvoer)
                _moet = _normaal(_verwacht)
            if not _ok:
                if _geval.get("aanroep"):
                    _info += "&lt;br&gt;verwachte returnwaarde: &lt;code&gt;" + _html(repr(_vwaarde)) + \
                             "&lt;/code&gt;&lt;br&gt;jouw returnwaarde: &lt;code&gt;" + _html(repr(_waarde)) + "&lt;/code&gt;"
                if not _gelijk(_uitvoer, _verwacht) or _normaal(_verwacht) != "":
                    _info += "&lt;br&gt;verwachte uitvoer:" + _blok(_normaal(_verwacht)) + \
                             "jouw uitvoer:" + _blok(_normaal(_uitvoer))
            self.appendResult(_ok, _jij, _moet, _info)


_Toets().main()

</tests>
      </program>
    </exercise>
  </section>
  <section xml:id="les-11c-s3">
    <title>Afstand in het platte vlak</title>
    <p>Punten in het platte vlak leggen we vast door twee zogenaamde coördinaten (<url href="http://en.wikipedia.org/wiki/Cartesian_coordinate_system" visual="en.wikipedia.org/wiki/Cartesian_coordinate_system">het Cartesisch coördinatensysteem</url>), de <em>x</em>-coördinaat en de <em>y</em>-coördinaat. We willen graag een functie schrijven waarvan de invoer een tweetal punten is, en waarbij de uitvoer de afstand tussen die twee punten is. Nu blijkt dat de functie <c>hypotenuse</c> ons helpt bij het uitvoeren van die taak!</p>
    <p>Meer in detail, stel dat het eerste punt coördinaten (x1, y1) heeft waarbij x1 en y1 reële getallen zijn, en stel dat het tweede punt coördinaten (x2, y2) heeft. Het sleutelidee is de rechthoekige driehoek te tekenen zoals in onderstaande tekening: de hypotenusa loopt van (x1, y1) naar (x2, y2) en de lengte ervan is gelijk aan de afstand tussen de twee punten.</p>
    <image source="images/11c-distance2d.png" width="67%"/>
    <p>Om de lengte van de hypotenusa te berekenen is het nodig de twee andere zijden van de driehoek te kennen. Dit kan door de definitie van van coördinaten te gebruiken (zie tekening): de horizontale verplaatsing is <em>a</em> = x1-x2 en de verticale verplaatsing is <em>b</em> = y1-y2.</p>
    <exercise label="oef-11c-distance">
      <title>Afstand in het platte vlak</title>
      <statement>
        <p>Neem aan dat <c>hypotenuse(a, b)</c> reeds is gedefinieerd. Definieer een functie <c>distance2D(x1, y1, x2, y2)</c> die de afstand tussen het punt (x1, y1) en het punt (x2, y2) berekent.</p>
      </statement>
      <program interactive="activecode" language="python" label="oef-11c-distance-code">
        <preamble visible="no">
import math
def hypotenuse(a, b):
    return math.sqrt(a * a + b * b)
</preamble>
        <code>

</code>
        <tests>
# Verborgen tests voor 11c.distance (gegenereerd door tools/harnas.py)
_OPLOSSING = 'def distance2D(x1, y1, x2, y2):\n    return hypotenuse(x2 - x1, y2 - y1)'
_GEVALLEN = [{'voorcode': '\nimport math\ndef hypotenuse(a, b):\n    return math.sqrt(a * a + b * b)\n', 'aanroep': 'distance2D(0, 0, 3, 4)'}, {'voorcode': '\nimport math\ndef hypotenuse(a, b):\n    return math.sqrt(a * a + b * b)\n', 'aanroep': 'distance2D(1, 2, 4, 6)'}, {'voorcode': '\nimport math\ndef hypotenuse(a, b):\n    return math.sqrt(a * a + b * b)\n', 'aanroep': 'distance2D(5, 5, 5, 5)'}, {'voorcode': '\nimport math\ndef hypotenuse(a, b):\n    return math.sqrt(a * a + b * b)\n', 'aanroep': 'distance2D(-1, -1, 2, 3)'}, {'voorcode': '\nimport math\ndef hypotenuse(a, b):\n    return math.sqrt(a * a + b * b)\n', 'aanroep': 'distance2D(0.5, 0, 0, 1.2)'}]
_VERGELIJK = 'getallen'
_TOLERANTIE = 1e-06
_HERGEBRUIK = False
_VERBODEN = []
_VERPLICHT = []
_MAXLENGTE = 0
_MAXWIJZIG = 0
_ORIGINEEL = ''

import re as _re
from unittest.gui import TestCaseGui


class _TeWeinigInvoer(Exception):
    pass


def _voer_uit(_bron, _geval, _student=False):
    _regels = [str(_r) for _r in _geval.get("invoer", [])]
    _uit = []

    def _print(*args, **kw):
        _sep = kw.get("sep", " ")
        _end = kw.get("end", "\n")
        if _sep is None:
            _sep = " "
        if _end is None:
            _end = "\n"
        _uit.append(_sep.join([str(_a) for _a in args]) + _end)

    def _input(prompt=""):
        if len(_regels) == 0:
            raise _TeWeinigInvoer("het programma vraagt meer invoer (input) dan er is")
        return _regels.pop(0)

    _g = {"__name__": "__main__", "print": _print, "input": _input}
    _fout = None
    _waarde = None
    if _student and _HERGEBRUIK:
        # Het programma van de leerling is net al uitgevoerd; gebruik die variabelen.
        try:
            _waarde = eval(_geval["aanroep"], globals())
        except Exception as _e:
            _fout = type(_e).__name__ + ": " + str(_e)
        return "", _waarde, _fout
    try:
        if _geval.get("voorcode"):
            exec(_geval["voorcode"], _g)
        exec(_bron, _g)
        if _geval.get("aanroep"):
            _waarde = eval(_geval["aanroep"], _g)
    except Exception as _e:
        _fout = type(_e).__name__ + ": " + str(_e)
    return "".join(_uit), _waarde, _fout


def _normaal(_s):
    # Spaties aan het eind van een regel tellen niet mee, en het (onzichtbare) DEL-teken ook niet.
    _r = [_x.replace(chr(127), "").rstrip() for _x in _s.split("\n")]
    while len(_r) &gt; 0 and _r[-1] == "":
        _r.pop()
    return "\n".join(_r)


def _getal(_s):
    try:
        return float(_s)
    except Exception:
        return None


def _stukken(_s):
    # Splits in getallen en overige tekens, zodat bv. "-13.333C" -&gt; ["-13.333", "C"].
    return _re.findall(r"[-+]?(?:[0-9]+[.]?[0-9]*|[.][0-9]+)(?:[eE][-+]?[0-9]+)?|[^\s0-9.+-]+|\S", _s)


def _gelijk(_a, _b):
    if _VERGELIJK == "woorden":
        return _a.split() == _b.split()
    if _VERGELIJK == "getallen":
        _wa = _stukken(_a)
        _wb = _stukken(_b)
        if len(_wa) != len(_wb):
            return False
        for _i in range(len(_wa)):
            _x = _getal(_wa[_i])
            _y = _getal(_wb[_i])
            if _x is not None and _y is not None:
                if abs(_x - _y) &gt; _TOLERANTIE * max(1.0, abs(_y)):
                    return False
            elif _wa[_i] != _wb[_i]:
                return False
        return True
    return _normaal(_a) == _normaal(_b)


def _afstand(_a, _b):
    """Bewerkingsafstand (Levenshtein) tussen twee strings."""
    _vorige = list(range(len(_b) + 1))
    for _i in range(len(_a)):
        _huidig = [_i + 1]
        for _j in range(len(_b)):
            _kosten = 0 if _a[_i] == _b[_j] else 1
            _huidig.append(min(_vorige[_j + 1] + 1, _huidig[_j] + 1, _vorige[_j] + _kosten))
        _vorige = _huidig
    return _vorige[-1]


def _zelfde(_a, _b):
    """Vergelijk returnwaarden; bij vergelijk="getallen" mogen kommagetallen iets afwijken."""
    if _VERGELIJK == "getallen" and type(_a) in (int, float) and type(_b) in (int, float) \
            and type(_a) != bool and type(_b) != bool:
        return abs(_a - _b) &lt;= _TOLERANTIE * max(1.0, abs(_b))
    if type(_a) in (list, tuple) and type(_b) == type(_a):
        if len(_a) != len(_b):
            return False
        for _i in range(len(_a)):
            if not _zelfde(_a[_i], _b[_i]):
                return False
        return True
    return type(_a) == type(_b) and _a == _b


def _html(_s):
    return str(_s).replace("&amp;", "&amp;amp;").replace("&lt;", "&amp;lt;").replace("&gt;", "&amp;gt;")


def _blok(_s):
    if _s == "":
        return "&lt;i&gt;(geen uitvoer)&lt;/i&gt;"
    return "&lt;pre style='text-align:left;margin:2px 0'&gt;" + _html(_s) + "&lt;/pre&gt;"


class _Toets(TestCaseGui):
    def test_1_programma(self):
        _bron = self.getEditorText()
        for _tekst, _melding in _VERPLICHT:
            self.assertTrue(_tekst in _bron, _melding)
        for _tekst, _melding in _VERBODEN:
            self.assertFalse(_tekst in _bron, _melding)
        if _MAXLENGTE:
            _lengte = len(_bron.rstrip())
            self.assertTrue(_lengte &lt;= _MAXLENGTE, "Je code is %d tekens lang; het mag hoogstens %d zijn."
                            % (_lengte, _MAXLENGTE))
        if _MAXWIJZIG:
            _d = _afstand(_normaal(_ORIGINEEL), _normaal(_bron))
            self.assertTrue(_d &lt;= _MAXWIJZIG, "Je hebt %d teken(s) veranderd; het mogen er hoogstens %d zijn."
                            % (_d, _MAXWIJZIG))
        _nr = 0
        for _geval in _GEVALLEN:
            _nr += 1
            _verwacht, _vwaarde, _vfout = _voer_uit(_OPLOSSING, _geval)
            _uitvoer, _waarde, _fout = _voer_uit(_bron, _geval, True)
            _info = "&lt;b&gt;Test %d&lt;/b&gt;" % _nr
            if _geval.get("voorcode"):
                _info += "&lt;br&gt;vooraf: &lt;code&gt;" + _html("; ".join(_geval["voorcode"].strip().split("\n"))) + "&lt;/code&gt;"
            if _geval.get("invoer"):
                _info += "&lt;br&gt;invoer:" + _blok("\n".join([str(_r) for _r in _geval["invoer"]]))
            if _geval.get("aanroep"):
                _info += "&lt;br&gt;aanroep: &lt;code&gt;" + _html(_geval["aanroep"]) + "&lt;/code&gt;"
            if _fout is not None and _vfout is None:
                self.appendResult(False, _fout, "", _info + "&lt;br&gt;fout: &lt;code&gt;" + _html(_fout) + "&lt;/code&gt;")
                continue
            _ok = _gelijk(_uitvoer, _verwacht)
            if _geval.get("aanroep"):
                _ok = _ok and _zelfde(_waarde, _vwaarde)
                _jij = _waarde
                _moet = _vwaarde
                if _normaal(_verwacht) != "" or _normaal(_uitvoer) != "":
                    _jij = _normaal(_uitvoer) + " / " + repr(_waarde)
                    _moet = _normaal(_verwacht) + " / " + repr(_vwaarde)
            else:
                _jij = _normaal(_uitvoer)
                _moet = _normaal(_verwacht)
            if not _ok:
                if _geval.get("aanroep"):
                    _info += "&lt;br&gt;verwachte returnwaarde: &lt;code&gt;" + _html(repr(_vwaarde)) + \
                             "&lt;/code&gt;&lt;br&gt;jouw returnwaarde: &lt;code&gt;" + _html(repr(_waarde)) + "&lt;/code&gt;"
                if not _gelijk(_uitvoer, _verwacht) or _normaal(_verwacht) != "":
                    _info += "&lt;br&gt;verwachte uitvoer:" + _blok(_normaal(_verwacht)) + \
                             "jouw uitvoer:" + _blok(_normaal(_uitvoer))
            self.appendResult(_ok, _jij, _moet, _info)


_Toets().main()

</tests>
      </program>
    </exercise>
  </section>
  <section xml:id="les-11c-s4">
    <title>Omtrek van een willekeurige driehoek</title>
    <p>We gaan nu de laatste oefening van deze les maken. Bedenk dat de omtrek van een driehoek de som is van de lengten van de zijden; en merk op dat de lengte van een zijde gelijk is aan de afstand tussen de twee bijbehorende hoekpunten.</p>
    <exercise label="oef-11c-perimeter">
      <title>Secure the Perimeter</title>
      <statement>
        <p>Neem aan dat <c>distance2D(x1, y1, x2, y2)</c> reeds is gedefinieerd. Definieer een functie <c>trianglePerimeter(xA, yA, xB, yB, xC, yC)</c> die de omtrek berekent van een driehoek waarvan de hoekpunten zijn vastgelegd door (xA, yA), (xB, yB) and (xC, yC).</p>
      </statement>
      <program interactive="activecode" language="python" label="oef-11c-perimeter-code">
        <preamble visible="no">
import math
def hypotenuse(a, b):
    return math.sqrt(a * a + b * b)

def distance2D(x1, y1, x2, y2):
    return hypotenuse(x2 - x1, y2 - y1)
</preamble>
        <code>

</code>
        <tests>
# Verborgen tests voor 11c.perimeter (gegenereerd door tools/harnas.py)
_OPLOSSING = 'def trianglePerimeter(xA, yA, xB, yB, xC, yC):\n    return distance2D(xA, yA, xB, yB) + distance2D(xB, yB, xC, yC) + distance2D(xC, yC, xA, yA)'
_GEVALLEN = [{'voorcode': '\nimport math\ndef hypotenuse(a, b):\n    return math.sqrt(a * a + b * b)\n\ndef distance2D(x1, y1, x2, y2):\n    return hypotenuse(x2 - x1, y2 - y1)\n', 'aanroep': 'trianglePerimeter(0, 0, 3, 0, 0, 4)'}, {'voorcode': '\nimport math\ndef hypotenuse(a, b):\n    return math.sqrt(a * a + b * b)\n\ndef distance2D(x1, y1, x2, y2):\n    return hypotenuse(x2 - x1, y2 - y1)\n', 'aanroep': 'trianglePerimeter(1, 1, 1, 1, 1, 1)'}, {'voorcode': '\nimport math\ndef hypotenuse(a, b):\n    return math.sqrt(a * a + b * b)\n\ndef distance2D(x1, y1, x2, y2):\n    return hypotenuse(x2 - x1, y2 - y1)\n', 'aanroep': 'trianglePerimeter(-2, 0, 2, 0, 0, 5)'}, {'voorcode': '\nimport math\ndef hypotenuse(a, b):\n    return math.sqrt(a * a + b * b)\n\ndef distance2D(x1, y1, x2, y2):\n    return hypotenuse(x2 - x1, y2 - y1)\n', 'aanroep': 'trianglePerimeter(0, 0, 1, 0, 0.5, 0.866)'}]
_VERGELIJK = 'getallen'
_TOLERANTIE = 1e-06
_HERGEBRUIK = False
_VERBODEN = []
_VERPLICHT = []
_MAXLENGTE = 0
_MAXWIJZIG = 0
_ORIGINEEL = ''

import re as _re
from unittest.gui import TestCaseGui


class _TeWeinigInvoer(Exception):
    pass


def _voer_uit(_bron, _geval, _student=False):
    _regels = [str(_r) for _r in _geval.get("invoer", [])]
    _uit = []

    def _print(*args, **kw):
        _sep = kw.get("sep", " ")
        _end = kw.get("end", "\n")
        if _sep is None:
            _sep = " "
        if _end is None:
            _end = "\n"
        _uit.append(_sep.join([str(_a) for _a in args]) + _end)

    def _input(prompt=""):
        if len(_regels) == 0:
            raise _TeWeinigInvoer("het programma vraagt meer invoer (input) dan er is")
        return _regels.pop(0)

    _g = {"__name__": "__main__", "print": _print, "input": _input}
    _fout = None
    _waarde = None
    if _student and _HERGEBRUIK:
        # Het programma van de leerling is net al uitgevoerd; gebruik die variabelen.
        try:
            _waarde = eval(_geval["aanroep"], globals())
        except Exception as _e:
            _fout = type(_e).__name__ + ": " + str(_e)
        return "", _waarde, _fout
    try:
        if _geval.get("voorcode"):
            exec(_geval["voorcode"], _g)
        exec(_bron, _g)
        if _geval.get("aanroep"):
            _waarde = eval(_geval["aanroep"], _g)
    except Exception as _e:
        _fout = type(_e).__name__ + ": " + str(_e)
    return "".join(_uit), _waarde, _fout


def _normaal(_s):
    # Spaties aan het eind van een regel tellen niet mee, en het (onzichtbare) DEL-teken ook niet.
    _r = [_x.replace(chr(127), "").rstrip() for _x in _s.split("\n")]
    while len(_r) &gt; 0 and _r[-1] == "":
        _r.pop()
    return "\n".join(_r)


def _getal(_s):
    try:
        return float(_s)
    except Exception:
        return None


def _stukken(_s):
    # Splits in getallen en overige tekens, zodat bv. "-13.333C" -&gt; ["-13.333", "C"].
    return _re.findall(r"[-+]?(?:[0-9]+[.]?[0-9]*|[.][0-9]+)(?:[eE][-+]?[0-9]+)?|[^\s0-9.+-]+|\S", _s)


def _gelijk(_a, _b):
    if _VERGELIJK == "woorden":
        return _a.split() == _b.split()
    if _VERGELIJK == "getallen":
        _wa = _stukken(_a)
        _wb = _stukken(_b)
        if len(_wa) != len(_wb):
            return False
        for _i in range(len(_wa)):
            _x = _getal(_wa[_i])
            _y = _getal(_wb[_i])
            if _x is not None and _y is not None:
                if abs(_x - _y) &gt; _TOLERANTIE * max(1.0, abs(_y)):
                    return False
            elif _wa[_i] != _wb[_i]:
                return False
        return True
    return _normaal(_a) == _normaal(_b)


def _afstand(_a, _b):
    """Bewerkingsafstand (Levenshtein) tussen twee strings."""
    _vorige = list(range(len(_b) + 1))
    for _i in range(len(_a)):
        _huidig = [_i + 1]
        for _j in range(len(_b)):
            _kosten = 0 if _a[_i] == _b[_j] else 1
            _huidig.append(min(_vorige[_j + 1] + 1, _huidig[_j] + 1, _vorige[_j] + _kosten))
        _vorige = _huidig
    return _vorige[-1]


def _zelfde(_a, _b):
    """Vergelijk returnwaarden; bij vergelijk="getallen" mogen kommagetallen iets afwijken."""
    if _VERGELIJK == "getallen" and type(_a) in (int, float) and type(_b) in (int, float) \
            and type(_a) != bool and type(_b) != bool:
        return abs(_a - _b) &lt;= _TOLERANTIE * max(1.0, abs(_b))
    if type(_a) in (list, tuple) and type(_b) == type(_a):
        if len(_a) != len(_b):
            return False
        for _i in range(len(_a)):
            if not _zelfde(_a[_i], _b[_i]):
                return False
        return True
    return type(_a) == type(_b) and _a == _b


def _html(_s):
    return str(_s).replace("&amp;", "&amp;amp;").replace("&lt;", "&amp;lt;").replace("&gt;", "&amp;gt;")


def _blok(_s):
    if _s == "":
        return "&lt;i&gt;(geen uitvoer)&lt;/i&gt;"
    return "&lt;pre style='text-align:left;margin:2px 0'&gt;" + _html(_s) + "&lt;/pre&gt;"


class _Toets(TestCaseGui):
    def test_1_programma(self):
        _bron = self.getEditorText()
        for _tekst, _melding in _VERPLICHT:
            self.assertTrue(_tekst in _bron, _melding)
        for _tekst, _melding in _VERBODEN:
            self.assertFalse(_tekst in _bron, _melding)
        if _MAXLENGTE:
            _lengte = len(_bron.rstrip())
            self.assertTrue(_lengte &lt;= _MAXLENGTE, "Je code is %d tekens lang; het mag hoogstens %d zijn."
                            % (_lengte, _MAXLENGTE))
        if _MAXWIJZIG:
            _d = _afstand(_normaal(_ORIGINEEL), _normaal(_bron))
            self.assertTrue(_d &lt;= _MAXWIJZIG, "Je hebt %d teken(s) veranderd; het mogen er hoogstens %d zijn."
                            % (_d, _MAXWIJZIG))
        _nr = 0
        for _geval in _GEVALLEN:
            _nr += 1
            _verwacht, _vwaarde, _vfout = _voer_uit(_OPLOSSING, _geval)
            _uitvoer, _waarde, _fout = _voer_uit(_bron, _geval, True)
            _info = "&lt;b&gt;Test %d&lt;/b&gt;" % _nr
            if _geval.get("voorcode"):
                _info += "&lt;br&gt;vooraf: &lt;code&gt;" + _html("; ".join(_geval["voorcode"].strip().split("\n"))) + "&lt;/code&gt;"
            if _geval.get("invoer"):
                _info += "&lt;br&gt;invoer:" + _blok("\n".join([str(_r) for _r in _geval["invoer"]]))
            if _geval.get("aanroep"):
                _info += "&lt;br&gt;aanroep: &lt;code&gt;" + _html(_geval["aanroep"]) + "&lt;/code&gt;"
            if _fout is not None and _vfout is None:
                self.appendResult(False, _fout, "", _info + "&lt;br&gt;fout: &lt;code&gt;" + _html(_fout) + "&lt;/code&gt;")
                continue
            _ok = _gelijk(_uitvoer, _verwacht)
            if _geval.get("aanroep"):
                _ok = _ok and _zelfde(_waarde, _vwaarde)
                _jij = _waarde
                _moet = _vwaarde
                if _normaal(_verwacht) != "" or _normaal(_uitvoer) != "":
                    _jij = _normaal(_uitvoer) + " / " + repr(_waarde)
                    _moet = _normaal(_verwacht) + " / " + repr(_vwaarde)
            else:
                _jij = _normaal(_uitvoer)
                _moet = _normaal(_verwacht)
            if not _ok:
                if _geval.get("aanroep"):
                    _info += "&lt;br&gt;verwachte returnwaarde: &lt;code&gt;" + _html(repr(_vwaarde)) + \
                             "&lt;/code&gt;&lt;br&gt;jouw returnwaarde: &lt;code&gt;" + _html(repr(_waarde)) + "&lt;/code&gt;"
                if not _gelijk(_uitvoer, _verwacht) or _normaal(_verwacht) != "":
                    _info += "&lt;br&gt;verwachte uitvoer:" + _blok(_normaal(_verwacht)) + \
                             "jouw uitvoer:" + _blok(_normaal(_uitvoer))
            self.appendResult(_ok, _jij, _moet, _info)


_Toets().main()

</tests>
      </program>
    </exercise>
    <p>Je bent nu klaar om naar de volgende les te gaan!</p>
  </section>
</chapter>
