<?xml version="1.0" encoding="UTF-8"?>
<!-- Gegenereerd door tools/convert.py uit archief/ ; niet met de hand bewerken. -->
<chapter xml:id="les-02x">
  <title>2X: Extra oefening</title>
  <introduction>
    <p>Deze les, met extra oefeningen over de functies <c>min</c> en <c>max</c>, is optioneel. Als je bij een opgave vast komt te zitten, sla deze dan over en bekijk de volgende opgave. Bekijk de overgeslagen opgave later nog eens.</p>
  </introduction>
  <section xml:id="les-02x-s1">
    <title>Een complexe uitdrukking vereenvoudigen</title>
    <exercise label="oef-2x-simplify">
      <title>Vereenvoudigen</title>
      <statement>
        <p>Wat is de vereenvoudiging van de volgende uitdrukking?</p>
        <pre>
max(x - 3, min(x + 10, x + 5))
</pre>
      </statement>
      <choices randomize="yes">
        <choice correct="yes">
          <statement>
            <p>x + 5</p>
          </statement>
          <feedback>
            <p>Correct! Omdat <c>x + 5</c> altijd kleiner is dan <c>x + 10</c>, kunnen we <c>min(x + 10, x + 5)</c> vereenvoudigen tot alleen maar <c>x + 5</c>. Dan werken we verder en wordt de uitdrukking <c>max(x - 3, x + 5)</c>. Omdat <c>x + 5</c> altijd het maximum is van de twee getallen, is dat het resultaat.</p>
          </feedback>
        </choice>
        <choice>
          <statement>
            <p>x + 10</p>
          </statement>
        </choice>
        <choice>
          <statement>
            <p>x - 3</p>
          </statement>
        </choice>
      </choices>
    </exercise>
  </section>
  <section xml:id="les-02x-s2">
    <title>Een eenvoudige expressie complexer maken</title>
    <exercise label="oef-2x-simplifyagain">
      <title>Complication</title>
      <statement>
        <p>Neem aan dat het nakijkprogramma twee variabelen A en B voor je definieert. Schrijf een programma dat de waarde afdrukt van</p>
        <pre>
min(A, B)
</pre>
        <p>Maar er is een beperking: je mag in je programma niet de functie <c>min</c> gebruiken. In plaats daarvan moet je de functie <c>max</c> op een slimme manier gebruiken teneinde <c>min</c> na te bootsen.</p>
      </statement>
      <program interactive="activecode" language="python" label="oef-2x-simplifyagain-code">
        <preamble visible="no">
A = 3
B = 8
</preamble>
        <code>

</code>
        <tests>
# Verborgen tests voor 2X.simplifyagain (gegenereerd door tools/harnas.py)
_OPLOSSING = 'print(-max(-A, -B))'
_GEVALLEN = [{'voorcode': 'A = 3\nB = 8'}, {'voorcode': 'A = 8\nB = 3'}, {'voorcode': 'A = -5\nB = 2'}, {'voorcode': 'A = 7\nB = 7'}]
_VERGELIJK = 'exact'
_TOLERANTIE = 1e-06
_HERGEBRUIK = False
_VERBODEN = [('min', 'Je mag de functie &lt;code&gt;min&lt;/code&gt; niet gebruiken.')]
_VERPLICHT = []
_MAXLENGTE = 0
_MAXWIJZIG = 0
_ORIGINEEL = ''

import re as _re
from unittest.gui import TestCaseGui


class _TeWeinigInvoer(Exception):
    pass


def _voer_uit(_bron, _geval, _student=False):
    _regels = [str(_r) for _r in _geval.get("invoer", [])]
    _uit = []

    def _print(*args, **kw):
        _sep = kw.get("sep", " ")
        _end = kw.get("end", "\n")
        if _sep is None:
            _sep = " "
        if _end is None:
            _end = "\n"
        _uit.append(_sep.join([str(_a) for _a in args]) + _end)

    def _input(prompt=""):
        if len(_regels) == 0:
            raise _TeWeinigInvoer("het programma vraagt meer invoer (input) dan er is")
        return _regels.pop(0)

    _g = {"__name__": "__main__", "print": _print, "input": _input}
    _fout = None
    _waarde = None
    if _student and _HERGEBRUIK:
        # Het programma van de leerling is net al uitgevoerd; gebruik die variabelen.
        try:
            _waarde = eval(_geval["aanroep"], globals())
        except Exception as _e:
            _fout = type(_e).__name__ + ": " + str(_e)
        return "", _waarde, _fout
    try:
        if _geval.get("voorcode"):
            exec(_geval["voorcode"], _g)
        exec(_bron, _g)
        if _geval.get("aanroep"):
            _waarde = eval(_geval["aanroep"], _g)
    except Exception as _e:
        _fout = type(_e).__name__ + ": " + str(_e)
    return "".join(_uit), _waarde, _fout


def _normaal(_s):
    # Spaties aan het eind van een regel tellen niet mee, en het (onzichtbare) DEL-teken ook niet.
    _r = [_x.replace(chr(127), "").rstrip() for _x in _s.split("\n")]
    while len(_r) &gt; 0 and _r[-1] == "":
        _r.pop()
    return "\n".join(_r)


def _getal(_s):
    try:
        return float(_s)
    except Exception:
        return None


def _stukken(_s):
    # Splits in getallen en overige tekens, zodat bv. "-13.333C" -&gt; ["-13.333", "C"].
    return _re.findall(r"[-+]?(?:[0-9]+[.]?[0-9]*|[.][0-9]+)(?:[eE][-+]?[0-9]+)?|[^\s0-9.+-]+|\S", _s)


def _gelijk(_a, _b):
    if _VERGELIJK == "woorden":
        return _a.split() == _b.split()
    if _VERGELIJK == "getallen":
        _wa = _stukken(_a)
        _wb = _stukken(_b)
        if len(_wa) != len(_wb):
            return False
        for _i in range(len(_wa)):
            _x = _getal(_wa[_i])
            _y = _getal(_wb[_i])
            if _x is not None and _y is not None:
                if abs(_x - _y) &gt; _TOLERANTIE * max(1.0, abs(_y)):
                    return False
            elif _wa[_i] != _wb[_i]:
                return False
        return True
    return _normaal(_a) == _normaal(_b)


def _afstand(_a, _b):
    """Bewerkingsafstand (Levenshtein) tussen twee strings."""
    _vorige = list(range(len(_b) + 1))
    for _i in range(len(_a)):
        _huidig = [_i + 1]
        for _j in range(len(_b)):
            _kosten = 0 if _a[_i] == _b[_j] else 1
            _huidig.append(min(_vorige[_j + 1] + 1, _huidig[_j] + 1, _vorige[_j] + _kosten))
        _vorige = _huidig
    return _vorige[-1]


def _zelfde(_a, _b):
    """Vergelijk returnwaarden; bij vergelijk="getallen" mogen kommagetallen iets afwijken."""
    if _VERGELIJK == "getallen" and type(_a) in (int, float) and type(_b) in (int, float) \
            and type(_a) != bool and type(_b) != bool:
        return abs(_a - _b) &lt;= _TOLERANTIE * max(1.0, abs(_b))
    if type(_a) in (list, tuple) and type(_b) == type(_a):
        if len(_a) != len(_b):
            return False
        for _i in range(len(_a)):
            if not _zelfde(_a[_i], _b[_i]):
                return False
        return True
    return type(_a) == type(_b) and _a == _b


def _html(_s):
    return str(_s).replace("&amp;", "&amp;amp;").replace("&lt;", "&amp;lt;").replace("&gt;", "&amp;gt;")


def _blok(_s):
    if _s == "":
        return "&lt;i&gt;(geen uitvoer)&lt;/i&gt;"
    return "&lt;pre style='text-align:left;margin:2px 0'&gt;" + _html(_s) + "&lt;/pre&gt;"


class _Toets(TestCaseGui):
    def test_1_programma(self):
        _bron = self.getEditorText()
        for _tekst, _melding in _VERPLICHT:
            self.assertTrue(_tekst in _bron, _melding)
        for _tekst, _melding in _VERBODEN:
            self.assertFalse(_tekst in _bron, _melding)
        if _MAXLENGTE:
            _lengte = len(_bron.rstrip())
            self.assertTrue(_lengte &lt;= _MAXLENGTE, "Je code is %d tekens lang; het mag hoogstens %d zijn."
                            % (_lengte, _MAXLENGTE))
        if _MAXWIJZIG:
            _d = _afstand(_normaal(_ORIGINEEL), _normaal(_bron))
            self.assertTrue(_d &lt;= _MAXWIJZIG, "Je hebt %d teken(s) veranderd; het mogen er hoogstens %d zijn."
                            % (_d, _MAXWIJZIG))
        _nr = 0
        for _geval in _GEVALLEN:
            _nr += 1
            _verwacht, _vwaarde, _vfout = _voer_uit(_OPLOSSING, _geval)
            _uitvoer, _waarde, _fout = _voer_uit(_bron, _geval, True)
            _info = "&lt;b&gt;Test %d&lt;/b&gt;" % _nr
            if _geval.get("voorcode"):
                _info += "&lt;br&gt;vooraf: &lt;code&gt;" + _html("; ".join(_geval["voorcode"].strip().split("\n"))) + "&lt;/code&gt;"
            if _geval.get("invoer"):
                _info += "&lt;br&gt;invoer:" + _blok("\n".join([str(_r) for _r in _geval["invoer"]]))
            if _geval.get("aanroep"):
                _info += "&lt;br&gt;aanroep: &lt;code&gt;" + _html(_geval["aanroep"]) + "&lt;/code&gt;"
            if _fout is not None and _vfout is None:
                self.appendResult(False, _fout, "", _info + "&lt;br&gt;fout: &lt;code&gt;" + _html(_fout) + "&lt;/code&gt;")
                continue
            _ok = _gelijk(_uitvoer, _verwacht)
            if _geval.get("aanroep"):
                _ok = _ok and _zelfde(_waarde, _vwaarde)
                _jij = _waarde
                _moet = _vwaarde
                if _normaal(_verwacht) != "" or _normaal(_uitvoer) != "":
                    _jij = _normaal(_uitvoer) + " / " + repr(_waarde)
                    _moet = _normaal(_verwacht) + " / " + repr(_vwaarde)
            else:
                _jij = _normaal(_uitvoer)
                _moet = _normaal(_verwacht)
            if not _ok:
                if _geval.get("aanroep"):
                    _info += "&lt;br&gt;verwachte returnwaarde: &lt;code&gt;" + _html(repr(_vwaarde)) + \
                             "&lt;/code&gt;&lt;br&gt;jouw returnwaarde: &lt;code&gt;" + _html(repr(_waarde)) + "&lt;/code&gt;"
                if not _gelijk(_uitvoer, _verwacht) or _normaal(_verwacht) != "":
                    _info += "&lt;br&gt;verwachte uitvoer:" + _blok(_normaal(_verwacht)) + \
                             "jouw uitvoer:" + _blok(_normaal(_uitvoer))
            self.appendResult(_ok, _jij, _moet, _info)


_Toets().main()

</tests>
      </program>
      <hint>
        <title>Hint, Methode 1</title>
        <p>Wat is <c>max(-A, -B)</c>?</p>
      </hint>
      <hint>
        <title>Hint, Methode 2</title>
        <p>Wat is <c>min(A, B)+max(A, B)</c>?</p>
      </hint>
    </exercise>
  </section>
  <section xml:id="les-02x-s3">
    <title>Betaalrekenmachine</title>
    <exercise label="oef-2x-saldo">
      <title>Betaalrekenmachine</title>
      <statement>
        <p>Een credit card-bedrijf berekent de "minimale betaling" volgens de volgende regel: de minimale betaling is gelijk aan ofwel € 10 of 2,1% van het openstaand saldo van de klant, indien dat meer is, maar als dit hoger is dan het openstaand saldo, dan is de minimale betaling het openstaande saldo. Schrijf een programma dat de minimale betaling berekent door gebruik te maken van de functies <c>min</c> en <c>max</c>. Neem daarbij aan dat de variabele <c>saldo</c> het openstaand saldo van de klant weergeeft. Je programma hoeft niet het euroteken te printen.  Voorbeeld<em> 1</em>: als je <c>saldo</c> gelijk is aan <c>1000</c>, dan dient je programma <c>21</c> te printen.   Voorbeeld<em> 2</em>: als je <c>saldo</c> gelijk is aan <c>600</c>, dan dient je programma <c>12.6</c> te printen.   Voorbeeld<em> 3</em>: als je <c>saldo</c> gelijk is aan <c>25</c>, dan dient je programma <c>10</c> te printen.   Voorbeeld<em> 4</em>: als je <c>saldo</c> gelijk is aan <c>8</c>, dan dient je programma <c>8</c> te printen.</p>
      </statement>
      <program interactive="activecode" language="python" label="oef-2x-saldo-code">
        <preamble visible="no">
saldo = 1000
</preamble>
        <code>

</code>
        <tests>
# Verborgen tests voor 2X.saldo (gegenereerd door tools/harnas.py)
_OPLOSSING = 'print(min(saldo, max(10, 0.021 * saldo)))'
_GEVALLEN = [{'voorcode': 'saldo = 1000'}, {'voorcode': 'saldo = 600'}, {'voorcode': 'saldo = 25'}, {'voorcode': 'saldo = 8'}, {'voorcode': 'saldo = 476.19'}]
_VERGELIJK = 'getallen'
_TOLERANTIE = 1e-06
_HERGEBRUIK = False
_VERBODEN = []
_VERPLICHT = [('min', 'Gebruik de functie &lt;code&gt;min&lt;/code&gt;.'), ('max', 'Gebruik de functie &lt;code&gt;max&lt;/code&gt;.')]
_MAXLENGTE = 0
_MAXWIJZIG = 0
_ORIGINEEL = ''

import re as _re
from unittest.gui import TestCaseGui


class _TeWeinigInvoer(Exception):
    pass


def _voer_uit(_bron, _geval, _student=False):
    _regels = [str(_r) for _r in _geval.get("invoer", [])]
    _uit = []

    def _print(*args, **kw):
        _sep = kw.get("sep", " ")
        _end = kw.get("end", "\n")
        if _sep is None:
            _sep = " "
        if _end is None:
            _end = "\n"
        _uit.append(_sep.join([str(_a) for _a in args]) + _end)

    def _input(prompt=""):
        if len(_regels) == 0:
            raise _TeWeinigInvoer("het programma vraagt meer invoer (input) dan er is")
        return _regels.pop(0)

    _g = {"__name__": "__main__", "print": _print, "input": _input}
    _fout = None
    _waarde = None
    if _student and _HERGEBRUIK:
        # Het programma van de leerling is net al uitgevoerd; gebruik die variabelen.
        try:
            _waarde = eval(_geval["aanroep"], globals())
        except Exception as _e:
            _fout = type(_e).__name__ + ": " + str(_e)
        return "", _waarde, _fout
    try:
        if _geval.get("voorcode"):
            exec(_geval["voorcode"], _g)
        exec(_bron, _g)
        if _geval.get("aanroep"):
            _waarde = eval(_geval["aanroep"], _g)
    except Exception as _e:
        _fout = type(_e).__name__ + ": " + str(_e)
    return "".join(_uit), _waarde, _fout


def _normaal(_s):
    # Spaties aan het eind van een regel tellen niet mee, en het (onzichtbare) DEL-teken ook niet.
    _r = [_x.replace(chr(127), "").rstrip() for _x in _s.split("\n")]
    while len(_r) &gt; 0 and _r[-1] == "":
        _r.pop()
    return "\n".join(_r)


def _getal(_s):
    try:
        return float(_s)
    except Exception:
        return None


def _stukken(_s):
    # Splits in getallen en overige tekens, zodat bv. "-13.333C" -&gt; ["-13.333", "C"].
    return _re.findall(r"[-+]?(?:[0-9]+[.]?[0-9]*|[.][0-9]+)(?:[eE][-+]?[0-9]+)?|[^\s0-9.+-]+|\S", _s)


def _gelijk(_a, _b):
    if _VERGELIJK == "woorden":
        return _a.split() == _b.split()
    if _VERGELIJK == "getallen":
        _wa = _stukken(_a)
        _wb = _stukken(_b)
        if len(_wa) != len(_wb):
            return False
        for _i in range(len(_wa)):
            _x = _getal(_wa[_i])
            _y = _getal(_wb[_i])
            if _x is not None and _y is not None:
                if abs(_x - _y) &gt; _TOLERANTIE * max(1.0, abs(_y)):
                    return False
            elif _wa[_i] != _wb[_i]:
                return False
        return True
    return _normaal(_a) == _normaal(_b)


def _afstand(_a, _b):
    """Bewerkingsafstand (Levenshtein) tussen twee strings."""
    _vorige = list(range(len(_b) + 1))
    for _i in range(len(_a)):
        _huidig = [_i + 1]
        for _j in range(len(_b)):
            _kosten = 0 if _a[_i] == _b[_j] else 1
            _huidig.append(min(_vorige[_j + 1] + 1, _huidig[_j] + 1, _vorige[_j] + _kosten))
        _vorige = _huidig
    return _vorige[-1]


def _zelfde(_a, _b):
    """Vergelijk returnwaarden; bij vergelijk="getallen" mogen kommagetallen iets afwijken."""
    if _VERGELIJK == "getallen" and type(_a) in (int, float) and type(_b) in (int, float) \
            and type(_a) != bool and type(_b) != bool:
        return abs(_a - _b) &lt;= _TOLERANTIE * max(1.0, abs(_b))
    if type(_a) in (list, tuple) and type(_b) == type(_a):
        if len(_a) != len(_b):
            return False
        for _i in range(len(_a)):
            if not _zelfde(_a[_i], _b[_i]):
                return False
        return True
    return type(_a) == type(_b) and _a == _b


def _html(_s):
    return str(_s).replace("&amp;", "&amp;amp;").replace("&lt;", "&amp;lt;").replace("&gt;", "&amp;gt;")


def _blok(_s):
    if _s == "":
        return "&lt;i&gt;(geen uitvoer)&lt;/i&gt;"
    return "&lt;pre style='text-align:left;margin:2px 0'&gt;" + _html(_s) + "&lt;/pre&gt;"


class _Toets(TestCaseGui):
    def test_1_programma(self):
        _bron = self.getEditorText()
        for _tekst, _melding in _VERPLICHT:
            self.assertTrue(_tekst in _bron, _melding)
        for _tekst, _melding in _VERBODEN:
            self.assertFalse(_tekst in _bron, _melding)
        if _MAXLENGTE:
            _lengte = len(_bron.rstrip())
            self.assertTrue(_lengte &lt;= _MAXLENGTE, "Je code is %d tekens lang; het mag hoogstens %d zijn."
                            % (_lengte, _MAXLENGTE))
        if _MAXWIJZIG:
            _d = _afstand(_normaal(_ORIGINEEL), _normaal(_bron))
            self.assertTrue(_d &lt;= _MAXWIJZIG, "Je hebt %d teken(s) veranderd; het mogen er hoogstens %d zijn."
                            % (_d, _MAXWIJZIG))
        _nr = 0
        for _geval in _GEVALLEN:
            _nr += 1
            _verwacht, _vwaarde, _vfout = _voer_uit(_OPLOSSING, _geval)
            _uitvoer, _waarde, _fout = _voer_uit(_bron, _geval, True)
            _info = "&lt;b&gt;Test %d&lt;/b&gt;" % _nr
            if _geval.get("voorcode"):
                _info += "&lt;br&gt;vooraf: &lt;code&gt;" + _html("; ".join(_geval["voorcode"].strip().split("\n"))) + "&lt;/code&gt;"
            if _geval.get("invoer"):
                _info += "&lt;br&gt;invoer:" + _blok("\n".join([str(_r) for _r in _geval["invoer"]]))
            if _geval.get("aanroep"):
                _info += "&lt;br&gt;aanroep: &lt;code&gt;" + _html(_geval["aanroep"]) + "&lt;/code&gt;"
            if _fout is not None and _vfout is None:
                self.appendResult(False, _fout, "", _info + "&lt;br&gt;fout: &lt;code&gt;" + _html(_fout) + "&lt;/code&gt;")
                continue
            _ok = _gelijk(_uitvoer, _verwacht)
            if _geval.get("aanroep"):
                _ok = _ok and _zelfde(_waarde, _vwaarde)
                _jij = _waarde
                _moet = _vwaarde
                if _normaal(_verwacht) != "" or _normaal(_uitvoer) != "":
                    _jij = _normaal(_uitvoer) + " / " + repr(_waarde)
                    _moet = _normaal(_verwacht) + " / " + repr(_vwaarde)
            else:
                _jij = _normaal(_uitvoer)
                _moet = _normaal(_verwacht)
            if not _ok:
                if _geval.get("aanroep"):
                    _info += "&lt;br&gt;verwachte returnwaarde: &lt;code&gt;" + _html(repr(_vwaarde)) + \
                             "&lt;/code&gt;&lt;br&gt;jouw returnwaarde: &lt;code&gt;" + _html(repr(_waarde)) + "&lt;/code&gt;"
                if not _gelijk(_uitvoer, _verwacht) or _normaal(_verwacht) != "":
                    _info += "&lt;br&gt;verwachte uitvoer:" + _blok(_normaal(_verwacht)) + \
                             "jouw uitvoer:" + _blok(_normaal(_uitvoer))
            self.appendResult(_ok, _jij, _moet, _info)


_Toets().main()

</tests>
      </program>
      <hint>
        <p>Verdeel het programma in twee delen. Bereken eerst het maximum van <c>0.021</c> maal <c>saldo</c> en <c>10</c>. Dit is de gewenste betaling behalve wanneer het meer is dan het openstaand saldo. Gebruik nu <c>min</c> op de een of andere manier.</p>
      </hint>
    </exercise>
  </section>
  <section xml:id="les-02x-s4">
    <title>Versleep-oefening: sorteren</title>
    <p>De laatste opgave is een uitdagende opgave, over het sorteren van getallen maar wel op een vreemde manier. Er zijn betere, eenvoudigere en snellere <url href="http://en.wikipedia.org/wiki/Sorting_algorithm" visual="en.wikipedia.org/wiki/Sorting_algorithm">manieren om te sorteren </url>die je kunt leren nadat je onze lessen hebt doorgewerkt.</p>
    <p>Als je vast komt te zitten, sla dan deze opgave voorlopig over. Je kunt het later nog een keer doen. Bovendien kun je bijhouden wat je wel en niet hebt opgelost door te kijken op de pagina Mijn voortgang.</p>
    <exercise label="oef-2x-versleep-oefening-sorteren" language="python">
      <title>Versleep-oefening: Sorteren</title>
      <statement>
        <p><alert>Versleep-oefening: sorteren</alert>: zorg ervoor dat het programma de drie getallen x, y en z <alert>sorteert</alert> in opklimmende grootte waarbij <c>x</c> de kleinste waarde heeft, <c>y</c> de daaropvolgende waarde heeft, en <c>z</c> de grootste waarde heeft.  <em>Sleep met je muis de regels in de juiste volgorde. </em></p>
      </statement>
      <blocks>
        <block>
          <cline>tmp = max(x, y)</cline>
        </block>
        <block>
          <cline>x = min(x, y)</cline>
        </block>
        <block>
          <cline>y = tmp</cline>
        </block>
        <block>
          <cline>tmp = max(y, z)</cline>
        </block>
        <block>
          <cline>y = min(y, z)</cline>
        </block>
        <block>
          <cline>z = tmp</cline>
        </block>
        <block>
          <cline>tmp = max(x, y)</cline>
        </block>
        <block>
          <cline>x = min(x, y)</cline>
        </block>
        <block>
          <cline>y = tmp</cline>
        </block>
      </blocks>
      <hint>
        <p>Dit programma gebruikt een strategie waarbij eerst het paar (x, y), daarna het paar (y, z), en dan opnieuw (x, y) wordt bewerkt.</p>
      </hint>
    </exercise>
    <p>Ga verder met les 3!</p>
  </section>
</chapter>
