<?xml version="1.0" encoding="UTF-8"?>
<!-- Gegenereerd door tools/convert.py uit archief/ ; niet met de hand bewerken. -->
<chapter xml:id="les-03">
  <title>3: Commentaar en aanhalingstekens</title>
  <section xml:id="les-03-s0">
    <title>Inleiding</title>
    <p>Een computerprogramma beschrijft eenduidig de opeenvolgende elementaire berekeningen die we de computer willen laten doen. Voor een computer is dat voldoende; die stelt zich daar verder geen vragen bij. Een menselijke lezer van een computerprogramma wil echter ook het <em>waarom </em>van het programma begrijpen: hij of zij wilt begrijpen wat het <em>doel</em> is van elk stuk van het programma, wilt kunnen inzien dat dat stuk zijn doel inderdaad altijd correct bereikt, en wilt kunnen inzien hoe de verschillende stukken in elkaar passen en samen het overkoepelende doel van het programma realiseren. Bij het schrijven van een programma is het onze taak ervoor te zorgen dat het voor een lezer duidelijk is wat de antwoorden zijn op deze vragen. Dat doen we in de eerste plaats door middel van een doelmatige <em>naamgeving</em> van onze variabelen en functies, en een logische <em>structurering</em> van ons programma. Maar als dat niet volstaat, laat de programmeertaal ons toe <alert>commentaar</alert> in te voegen in ons programma. Dit zijn stukken tekst die door de computer genegeerd worden, en waarin we bijkomende verduidelijkingen kunnen geven bij ons programma.</p>
    <p>In Python is, in elke regel van een programma waar het symbool <c>#</c> (hekje) in voorkomt, het hekje het begin van commentaar. De rest van de regel wordt genegeerd wanneer het programma wordt uitgevoerd. Hier volgt een voorbeeld.</p>
    <example>
      <p>Het effect van commentaar.</p>
      <program interactive="activecode" language="python" label="les-03-vb-1">
        <code>
print(1) # hier volgt commentaar
# print(2)
print(3) # nog wat commentaar @!#!@$
</code>
      </program>
    </example>
    <p>Omdat de tweede regel start met het <c>#</c>-teken, zal Python die regel in zijn geheel negeren, en zoals je kunt zien werd nummer <c>2</c> niet afgedrukt. Normaal gebruik van commentaar omvat onder meer:<ul><li><p>uitleg over delen van het programma, voor jezelf en anderen om later nog eens door te nemen;</p></li><li><p>TODO-aantekeningen, die aangeven dat je later nog iets moet doen met betrekking tot die plaats in het programma;</p></li><li><p>een regel van het programma tijdelijk buiten werking zetten ("<em>uitcommentariëren</em>") zonder de regel te wissen. Je kunt de regel later dan weer gemakkelijk invoegen in het programma.</p></li></ul></p>
    <p>Hier volgt een oefening om dit te illustreren. Wanneer je teveel veranderd hebt en terug wilt keren naar de beginsituatie, klik dan op <alert>Reset code to default.</alert></p>
    <exercise label="oef-3-debug">
      <title>Seconden raden</title>
      <statement>
        <p>Verbeter (debug) dit programma zodanig dat het het aantal seconden in een week afdrukt.</p>
      </statement>
      <program interactive="activecode" language="python" label="oef-3-debug-code">
        <code>
# doel: druk het aantal seconden in een week af 
secondsPerMinute = 60
secondsPerHour = secondsPerMinute * 50 # TODO: controleer dit!
secondsPerDay = secondsPerHour * 24
daysPerWeek = 5
# daysPerWeek = daysPerWeek + 2 # weekenden tellen niet mee!?
print(secondsPerDay * daysPerWeek)
</code>
        <tests>
# Verborgen tests voor 3.debug (gegenereerd door tools/harnas.py)
_OPLOSSING = '# doel: druk het aantal seconden in een week af\nsecondsPerMinute = 60\nsecondsPerHour = secondsPerMinute * 60\nsecondsPerDay = secondsPerHour * 24\ndaysPerWeek = 7\nprint(secondsPerDay * daysPerWeek)'
_GEVALLEN = [{}]
_VERGELIJK = 'exact'
_TOLERANTIE = 1e-06
_HERGEBRUIK = False
_VERBODEN = []
_VERPLICHT = []
_MAXLENGTE = 0
_MAXWIJZIG = 0
_ORIGINEEL = ''

import re as _re
from unittest.gui import TestCaseGui


class _TeWeinigInvoer(Exception):
    pass


def _voer_uit(_bron, _geval, _student=False):
    _regels = [str(_r) for _r in _geval.get("invoer", [])]
    _uit = []

    def _print(*args, **kw):
        _sep = kw.get("sep", " ")
        _end = kw.get("end", "\n")
        if _sep is None:
            _sep = " "
        if _end is None:
            _end = "\n"
        _uit.append(_sep.join([str(_a) for _a in args]) + _end)

    def _input(prompt=""):
        if len(_regels) == 0:
            raise _TeWeinigInvoer("het programma vraagt meer invoer (input) dan er is")
        return _regels.pop(0)

    _g = {"__name__": "__main__", "print": _print, "input": _input}
    _fout = None
    _waarde = None
    if _student and _HERGEBRUIK:
        # Het programma van de leerling is net al uitgevoerd; gebruik die variabelen.
        try:
            _waarde = eval(_geval["aanroep"], globals())
        except Exception as _e:
            _fout = type(_e).__name__ + ": " + str(_e)
        return "", _waarde, _fout
    try:
        if _geval.get("voorcode"):
            exec(_geval["voorcode"], _g)
        exec(_bron, _g)
        if _geval.get("aanroep"):
            _waarde = eval(_geval["aanroep"], _g)
    except Exception as _e:
        _fout = type(_e).__name__ + ": " + str(_e)
    return "".join(_uit), _waarde, _fout


def _normaal(_s):
    # Spaties aan het eind van een regel tellen niet mee, en het (onzichtbare) DEL-teken ook niet.
    _r = [_x.replace(chr(127), "").rstrip() for _x in _s.split("\n")]
    while len(_r) &gt; 0 and _r[-1] == "":
        _r.pop()
    return "\n".join(_r)


def _getal(_s):
    try:
        return float(_s)
    except Exception:
        return None


def _stukken(_s):
    # Splits in getallen en overige tekens, zodat bv. "-13.333C" -&gt; ["-13.333", "C"].
    return _re.findall(r"[-+]?(?:[0-9]+[.]?[0-9]*|[.][0-9]+)(?:[eE][-+]?[0-9]+)?|[^\s0-9.+-]+|\S", _s)


def _gelijk(_a, _b):
    if _VERGELIJK == "woorden":
        return _a.split() == _b.split()
    if _VERGELIJK == "getallen":
        _wa = _stukken(_a)
        _wb = _stukken(_b)
        if len(_wa) != len(_wb):
            return False
        for _i in range(len(_wa)):
            _x = _getal(_wa[_i])
            _y = _getal(_wb[_i])
            if _x is not None and _y is not None:
                if abs(_x - _y) &gt; _TOLERANTIE * max(1.0, abs(_y)):
                    return False
            elif _wa[_i] != _wb[_i]:
                return False
        return True
    return _normaal(_a) == _normaal(_b)


def _afstand(_a, _b):
    """Bewerkingsafstand (Levenshtein) tussen twee strings."""
    _vorige = list(range(len(_b) + 1))
    for _i in range(len(_a)):
        _huidig = [_i + 1]
        for _j in range(len(_b)):
            _kosten = 0 if _a[_i] == _b[_j] else 1
            _huidig.append(min(_vorige[_j + 1] + 1, _huidig[_j] + 1, _vorige[_j] + _kosten))
        _vorige = _huidig
    return _vorige[-1]


def _zelfde(_a, _b):
    """Vergelijk returnwaarden; bij vergelijk="getallen" mogen kommagetallen iets afwijken."""
    if _VERGELIJK == "getallen" and type(_a) in (int, float) and type(_b) in (int, float) \
            and type(_a) != bool and type(_b) != bool:
        return abs(_a - _b) &lt;= _TOLERANTIE * max(1.0, abs(_b))
    if type(_a) in (list, tuple) and type(_b) == type(_a):
        if len(_a) != len(_b):
            return False
        for _i in range(len(_a)):
            if not _zelfde(_a[_i], _b[_i]):
                return False
        return True
    return type(_a) == type(_b) and _a == _b


def _html(_s):
    return str(_s).replace("&amp;", "&amp;amp;").replace("&lt;", "&amp;lt;").replace("&gt;", "&amp;gt;")


def _blok(_s):
    if _s == "":
        return "&lt;i&gt;(geen uitvoer)&lt;/i&gt;"
    return "&lt;pre style='text-align:left;margin:2px 0'&gt;" + _html(_s) + "&lt;/pre&gt;"


class _Toets(TestCaseGui):
    def test_1_programma(self):
        _bron = self.getEditorText()
        for _tekst, _melding in _VERPLICHT:
            self.assertTrue(_tekst in _bron, _melding)
        for _tekst, _melding in _VERBODEN:
            self.assertFalse(_tekst in _bron, _melding)
        if _MAXLENGTE:
            _lengte = len(_bron.rstrip())
            self.assertTrue(_lengte &lt;= _MAXLENGTE, "Je code is %d tekens lang; het mag hoogstens %d zijn."
                            % (_lengte, _MAXLENGTE))
        if _MAXWIJZIG:
            _d = _afstand(_normaal(_ORIGINEEL), _normaal(_bron))
            self.assertTrue(_d &lt;= _MAXWIJZIG, "Je hebt %d teken(s) veranderd; het mogen er hoogstens %d zijn."
                            % (_d, _MAXWIJZIG))
        _nr = 0
        for _geval in _GEVALLEN:
            _nr += 1
            _verwacht, _vwaarde, _vfout = _voer_uit(_OPLOSSING, _geval)
            _uitvoer, _waarde, _fout = _voer_uit(_bron, _geval, True)
            _info = "&lt;b&gt;Test %d&lt;/b&gt;" % _nr
            if _geval.get("voorcode"):
                _info += "&lt;br&gt;vooraf: &lt;code&gt;" + _html("; ".join(_geval["voorcode"].strip().split("\n"))) + "&lt;/code&gt;"
            if _geval.get("invoer"):
                _info += "&lt;br&gt;invoer:" + _blok("\n".join([str(_r) for _r in _geval["invoer"]]))
            if _geval.get("aanroep"):
                _info += "&lt;br&gt;aanroep: &lt;code&gt;" + _html(_geval["aanroep"]) + "&lt;/code&gt;"
            if _fout is not None and _vfout is None:
                self.appendResult(False, _fout, "", _info + "&lt;br&gt;fout: &lt;code&gt;" + _html(_fout) + "&lt;/code&gt;")
                continue
            _ok = _gelijk(_uitvoer, _verwacht)
            if _geval.get("aanroep"):
                _ok = _ok and _zelfde(_waarde, _vwaarde)
                _jij = _waarde
                _moet = _vwaarde
                if _normaal(_verwacht) != "" or _normaal(_uitvoer) != "":
                    _jij = _normaal(_uitvoer) + " / " + repr(_waarde)
                    _moet = _normaal(_verwacht) + " / " + repr(_vwaarde)
            else:
                _jij = _normaal(_uitvoer)
                _moet = _normaal(_verwacht)
            if not _ok:
                if _geval.get("aanroep"):
                    _info += "&lt;br&gt;verwachte returnwaarde: &lt;code&gt;" + _html(repr(_vwaarde)) + \
                             "&lt;/code&gt;&lt;br&gt;jouw returnwaarde: &lt;code&gt;" + _html(repr(_waarde)) + "&lt;/code&gt;"
                if not _gelijk(_uitvoer, _verwacht) or _normaal(_verwacht) != "":
                    _info += "&lt;br&gt;verwachte uitvoer:" + _blok(_normaal(_verwacht)) + \
                             "jouw uitvoer:" + _blok(_normaal(_uitvoer))
            self.appendResult(_ok, _jij, _moet, _info)


_Toets().main()

</tests>
      </program>
    </exercise>
  </section>
  <section xml:id="les-03-s1">
    <title>Tekenreeksen</title>
    <p>Tekenreeksen (Eng.: <alert><em>strings</em></alert>) zijn rijen van tekens (letters, cijfers, spaties, leestekens, enz., ook <em>karakters</em> genoemd), of met andere woorden stukjes tekst. Je kan een Python-programma in zijn geheel zien als één grote tekenreeks. Als je echter een rij tekens in een programma wilt laten behandelen door dat programma als een <em>waarde</em> die je bv. wilt toekennen aan een variabele of meegeven als argument aan een functie-oproep, dan moet je de tekens omsluiten met aanhalingstekens. Zo'n uitdrukking noemen we een <em>tekenreeks-citaat</em> (Eng.: <em>string literal</em>); ze "citeert" een tekenreeks. Bijvoorbeeld: in Les 0 was <c>"Hello, World!"</c> van ons allereerste programma een tekenreeks-citaat.</p>
    <p>Wanneer het hekje (<c>#</c>) verschijnt in een tekenreeks-citaat, dan wordt het niet behandeld als commentaar:</p>
    <example>
      <p>Een tekenreeks-citaat dat het <c>#-</c>teken bevat.</p>
      <program interactive="activecode" language="python" label="les-03-vb-2">
        <code>
myStr = "Om een boodschap achter te laten druk je op de #-toets van je telefoon."
print(myStr)
</code>
      </program>
    </example>
    <p>Precies zo zal <c>print("3 + 4")</c> niet het getal 7 afdrukken, maar alleen maar de tekenreeks <c>3 + 4</c>.</p>
  </section>
  <section xml:id="les-03-s2">
    <title>Ontsnapsequenties</title>
    <p>Wat als je de aanhalingstekens (<c>"</c>) in een tekenreeks wilt opnemen? Wanneer je <c>print("Ik zei "Wow!" tegen hem.")</c> wilt uitvoeren, zal dat een foutmelding tot gevolg hebben: het probleem is dat Python het citaat <c>"Ik zei "</c> ziet, gevolgd door <c>Wow!</c>, wat niet tot het citaat behoort en niets betekent in Python. Dat is niet wat we willen!</p>
    <program interactive="activecode" language="python" label="les-03-console-3">
      <code>
print("Ik zei "Wow!" tegen hem.")
</code>
    </program>
    <p>Er zijn in Python twee eenvoudige manieren om aanhalingstekens in een tekenreeks op te nemen.</p>
    <p><ul><li><p>Het is toegestaan om een citaat te laten beginnen en eindigen met een enkelvoudig aanhalingsteken (ook wel apostrof genoemd), zoals in <c>'bla bla'</c>. Dan is het toegestaan, om in het citaat dubbele aanhalingstekens op te nemen, zoals in <c>'Ik zei "Wow!" tegen hem.'</c></p><program interactive="activecode" language="python" label="les-03-console-4"><code>
print('Ik zei "Wow!" tegen hem.')
</code></program></li><li><p>Je kunt een <em>backslash</em> (<c>\</c>) gebruiken gevolgd door een aanhalingsteken (<c>\"</c> of <c>\'</c>). Dit wordt een <em>ontsnapsequentie</em> (Eng.: <em>escape sequence</em>) genoemd, omdat je zo het aanhalingsteken kunt laten ontsnappen aan haar gewoonlijke betekenis. We zeggen dat we het aanhalingsteken <em>escapen</em>. Python zal de backslash weghalen, en alleen het aanhalingsteken opnemen in de tekenreeks. Hier volgt een voorbeeld.</p><p>De backslashes beschermen de aanhalingstekens, maar worden niet opgenomen in de tekenreeks.</p><program interactive="activecode" language="python" label="les-03-vb-5"><code>
print('Er staan auto\'s op mijn brooddoos.')
print("We \"escapen\" hier de aanhalingstekens.")
</code></program></li></ul></p>
    <p>De aanhalingstekens zijn niet de enige tekens die je moet <em>escapen</em> in een citaat om ze letterlijk in de tekenreeks op te nemen. Dat geldt ook voor de backslash zelf: als je die wilt opnemen in een tekenreeks moet je haar <em>escapen</em> door middel van een tweede backslash:</p>
    <exercise label="oef-3-escapemulti">
      <title>Ontsnapsequenties</title>
      <statement>
        <p>Wat is de uitvoer van de opdracht <c>print("Backslashes \\ en apostroffen \' en aanhalingstekens \" en hekjes # zijn geweldig!")</c></p>
      </statement>
      <choices randomize="yes">
        <choice>
          <statement>
            <p>Geeft een fout.</p>
          </statement>
        </choice>
        <choice correct="yes">
          <statement>
            <p>Backslashes \ en apostroffen ' en aanhalingstekens " en hekjes # zijn geweldig!</p>
          </statement>
          <feedback>
            <p>Correct!</p>
          </feedback>
        </choice>
        <choice>
          <statement>
            <p>en hekjes</p>
          </statement>
        </choice>
        <choice>
          <statement>
            <p>Backslashes \\ en apostroffen \' en aanhalingstekens \" en hekjes # zijn geweldig!</p>
          </statement>
        </choice>
      </choices>
    </exercise>
    <exercise label="oef-3-escape">
      <title>The Great Escape</title>
      <statement>
        <p>Schrijf een programma dat afdrukt:</p>
        <pre>
Escape 't aanhalingsteken met 'n backslash: \"
</pre>
        <p>Opmerking: er is (zoals gewoonlijk) meer dan één manier om dit te doen! Bijvoorbeeld kun je de tekenreeks insluiten in enkelvoudige aanhalingstekens, of in dubbele aanhalingstekens. In beide gevallen is gebruik van een ontsnapsequentie noodzakelijk. Los, als extra oefening, het probleem op de ene manier op en daarna op de andere. Kun je het ook nog  op een derde manier oplossen?</p>
      </statement>
      <program interactive="activecode" language="python" label="oef-3-escape-code">
        <code>

</code>
        <tests>
# Verborgen tests voor 3.escape (gegenereerd door tools/harnas.py)
_OPLOSSING = 'print(\'Escape \\\'t aanhalingsteken met \\\'n backslash: \\\\"\')'
_GEVALLEN = [{}]
_VERGELIJK = 'exact'
_TOLERANTIE = 1e-06
_HERGEBRUIK = False
_VERBODEN = []
_VERPLICHT = []
_MAXLENGTE = 0
_MAXWIJZIG = 0
_ORIGINEEL = ''

import re as _re
from unittest.gui import TestCaseGui


class _TeWeinigInvoer(Exception):
    pass


def _voer_uit(_bron, _geval, _student=False):
    _regels = [str(_r) for _r in _geval.get("invoer", [])]
    _uit = []

    def _print(*args, **kw):
        _sep = kw.get("sep", " ")
        _end = kw.get("end", "\n")
        if _sep is None:
            _sep = " "
        if _end is None:
            _end = "\n"
        _uit.append(_sep.join([str(_a) for _a in args]) + _end)

    def _input(prompt=""):
        if len(_regels) == 0:
            raise _TeWeinigInvoer("het programma vraagt meer invoer (input) dan er is")
        return _regels.pop(0)

    _g = {"__name__": "__main__", "print": _print, "input": _input}
    _fout = None
    _waarde = None
    if _student and _HERGEBRUIK:
        # Het programma van de leerling is net al uitgevoerd; gebruik die variabelen.
        try:
            _waarde = eval(_geval["aanroep"], globals())
        except Exception as _e:
            _fout = type(_e).__name__ + ": " + str(_e)
        return "", _waarde, _fout
    try:
        if _geval.get("voorcode"):
            exec(_geval["voorcode"], _g)
        exec(_bron, _g)
        if _geval.get("aanroep"):
            _waarde = eval(_geval["aanroep"], _g)
    except Exception as _e:
        _fout = type(_e).__name__ + ": " + str(_e)
    return "".join(_uit), _waarde, _fout


def _normaal(_s):
    # Spaties aan het eind van een regel tellen niet mee, en het (onzichtbare) DEL-teken ook niet.
    _r = [_x.replace(chr(127), "").rstrip() for _x in _s.split("\n")]
    while len(_r) &gt; 0 and _r[-1] == "":
        _r.pop()
    return "\n".join(_r)


def _getal(_s):
    try:
        return float(_s)
    except Exception:
        return None


def _stukken(_s):
    # Splits in getallen en overige tekens, zodat bv. "-13.333C" -&gt; ["-13.333", "C"].
    return _re.findall(r"[-+]?(?:[0-9]+[.]?[0-9]*|[.][0-9]+)(?:[eE][-+]?[0-9]+)?|[^\s0-9.+-]+|\S", _s)


def _gelijk(_a, _b):
    if _VERGELIJK == "woorden":
        return _a.split() == _b.split()
    if _VERGELIJK == "getallen":
        _wa = _stukken(_a)
        _wb = _stukken(_b)
        if len(_wa) != len(_wb):
            return False
        for _i in range(len(_wa)):
            _x = _getal(_wa[_i])
            _y = _getal(_wb[_i])
            if _x is not None and _y is not None:
                if abs(_x - _y) &gt; _TOLERANTIE * max(1.0, abs(_y)):
                    return False
            elif _wa[_i] != _wb[_i]:
                return False
        return True
    return _normaal(_a) == _normaal(_b)


def _afstand(_a, _b):
    """Bewerkingsafstand (Levenshtein) tussen twee strings."""
    _vorige = list(range(len(_b) + 1))
    for _i in range(len(_a)):
        _huidig = [_i + 1]
        for _j in range(len(_b)):
            _kosten = 0 if _a[_i] == _b[_j] else 1
            _huidig.append(min(_vorige[_j + 1] + 1, _huidig[_j] + 1, _vorige[_j] + _kosten))
        _vorige = _huidig
    return _vorige[-1]


def _zelfde(_a, _b):
    """Vergelijk returnwaarden; bij vergelijk="getallen" mogen kommagetallen iets afwijken."""
    if _VERGELIJK == "getallen" and type(_a) in (int, float) and type(_b) in (int, float) \
            and type(_a) != bool and type(_b) != bool:
        return abs(_a - _b) &lt;= _TOLERANTIE * max(1.0, abs(_b))
    if type(_a) in (list, tuple) and type(_b) == type(_a):
        if len(_a) != len(_b):
            return False
        for _i in range(len(_a)):
            if not _zelfde(_a[_i], _b[_i]):
                return False
        return True
    return type(_a) == type(_b) and _a == _b


def _html(_s):
    return str(_s).replace("&amp;", "&amp;amp;").replace("&lt;", "&amp;lt;").replace("&gt;", "&amp;gt;")


def _blok(_s):
    if _s == "":
        return "&lt;i&gt;(geen uitvoer)&lt;/i&gt;"
    return "&lt;pre style='text-align:left;margin:2px 0'&gt;" + _html(_s) + "&lt;/pre&gt;"


class _Toets(TestCaseGui):
    def test_1_programma(self):
        _bron = self.getEditorText()
        for _tekst, _melding in _VERPLICHT:
            self.assertTrue(_tekst in _bron, _melding)
        for _tekst, _melding in _VERBODEN:
            self.assertFalse(_tekst in _bron, _melding)
        if _MAXLENGTE:
            _lengte = len(_bron.rstrip())
            self.assertTrue(_lengte &lt;= _MAXLENGTE, "Je code is %d tekens lang; het mag hoogstens %d zijn."
                            % (_lengte, _MAXLENGTE))
        if _MAXWIJZIG:
            _d = _afstand(_normaal(_ORIGINEEL), _normaal(_bron))
            self.assertTrue(_d &lt;= _MAXWIJZIG, "Je hebt %d teken(s) veranderd; het mogen er hoogstens %d zijn."
                            % (_d, _MAXWIJZIG))
        _nr = 0
        for _geval in _GEVALLEN:
            _nr += 1
            _verwacht, _vwaarde, _vfout = _voer_uit(_OPLOSSING, _geval)
            _uitvoer, _waarde, _fout = _voer_uit(_bron, _geval, True)
            _info = "&lt;b&gt;Test %d&lt;/b&gt;" % _nr
            if _geval.get("voorcode"):
                _info += "&lt;br&gt;vooraf: &lt;code&gt;" + _html("; ".join(_geval["voorcode"].strip().split("\n"))) + "&lt;/code&gt;"
            if _geval.get("invoer"):
                _info += "&lt;br&gt;invoer:" + _blok("\n".join([str(_r) for _r in _geval["invoer"]]))
            if _geval.get("aanroep"):
                _info += "&lt;br&gt;aanroep: &lt;code&gt;" + _html(_geval["aanroep"]) + "&lt;/code&gt;"
            if _fout is not None and _vfout is None:
                self.appendResult(False, _fout, "", _info + "&lt;br&gt;fout: &lt;code&gt;" + _html(_fout) + "&lt;/code&gt;")
                continue
            _ok = _gelijk(_uitvoer, _verwacht)
            if _geval.get("aanroep"):
                _ok = _ok and _zelfde(_waarde, _vwaarde)
                _jij = _waarde
                _moet = _vwaarde
                if _normaal(_verwacht) != "" or _normaal(_uitvoer) != "":
                    _jij = _normaal(_uitvoer) + " / " + repr(_waarde)
                    _moet = _normaal(_verwacht) + " / " + repr(_vwaarde)
            else:
                _jij = _normaal(_uitvoer)
                _moet = _normaal(_verwacht)
            if not _ok:
                if _geval.get("aanroep"):
                    _info += "&lt;br&gt;verwachte returnwaarde: &lt;code&gt;" + _html(repr(_vwaarde)) + \
                             "&lt;/code&gt;&lt;br&gt;jouw returnwaarde: &lt;code&gt;" + _html(repr(_waarde)) + "&lt;/code&gt;"
                if not _gelijk(_uitvoer, _verwacht) or _normaal(_verwacht) != "":
                    _info += "&lt;br&gt;verwachte uitvoer:" + _blok(_normaal(_verwacht)) + \
                             "jouw uitvoer:" + _blok(_normaal(_uitvoer))
            self.appendResult(_ok, _jij, _moet, _info)


_Toets().main()

</tests>
      </program>
      <hint>
        <p>Bedenk dat <c>print</c> meerdere argumenten kan hebben.</p>
      </hint>
    </exercise>
    <p>Er zijn nog meer ontsnapsequenties, zoals om een regeleinde in de tekenreeks op te nemen; die zullen we nu niet behandelen. Op naar de volgende les!</p>
  </section>
</chapter>
