<?xml version="1.0" encoding="UTF-8"?>
<!-- Gegenereerd door tools/convert.py uit archief/ ; niet met de hand bewerken. -->
<chapter xml:id="les-10">
  <title>10: Functies definiëren</title>
  <introduction>
    <p>In deze les laten we je het belangrijkste idee van programmeren zien: je eigen functies definiëren! Functies zorgen ervoor dat je programma korter wordt, beter gestructureerd, en in onderdelen herbruikbaar. In latere lessen zullen we andere voordelen zien zoals <em>recursie</em>.</p>
    <p><alert>Een functie definiëren</alert></p>
    <p>We beginnen met een voorbeeld. Ken je de functie <c>abs</c> nog? Ze neemt één argument (een getal <em>x</em>), en ze geeft als resultaat de absolute waarde van <em>x</em> terug (m.a.w. <em>x</em> wanneer <em>x </em>≥ 0, en -<em>x</em> wanneer <em>x</em> &lt; 0). Kijk hoe we deze functie in Python definiëren:</p>
    <pre>
def abs(x):     # definieer een functie met als naam 'abs' die één parameter heeft namelijk 'x'
  if x &gt;= 0:    # het lichaam van de functie start hier
    return x
  else:
    return -x   # het lichaam van de functie stopt hier
</pre>
    <p>(Klik hier om deze code in de console te testen.)</p>
    <program interactive="activecode" language="python" label="les-10-console-1">
      <code>
def abs(x): #Functiedefinitie
  if x &amp;gt;= 0:
    return x
  else:  
    return -x

# Het hoofd-programma begint hier
print(abs(-5))
print(abs(10))
</code>
    </program>
  </introduction>
  <section xml:id="les-10-s1">
    <title>De twee delen van een functiedefinitie</title>
    <p>In de eerste regel van een functiedefinitie staan de <em>naam</em> van de functie en haar <em>parameters</em>. De regel is van de vorm</p>
    <p><c>def «functie-naam»(«parameterlijst»):</c></p>
    <p>In het voorbeeld hierboven is <c>abs</c> de naam van de functie en bestaat de parameterlijst uit één element, namelijk <c>x</c>. Het wordt een <em>parameterlijst</em> genoemd omdat het meerdere parameters kunnen zijn, zoals <c>x, y, z</c>; het is ook mogelijk om 0 parameters te hebben, met andere woorden een lege lijst. Bij een oproep van een functie moet je precies zoveel argumenten opgeven als de functie parameters heeft, dus voor elke parameter van de functie één argument. Bij een oproep <c>abs(5)</c> van de hierboven gedefinieerde functie <c>abs</c> zeggen we dan ook dat 5 <em>het argument</em> is <em>voor de parameter</em> <c>x</c>.</p>
    <p>Alles na de eerste regel van de definitie van een functie is het <em>lichaam</em> van de functie. Het lichaam is een ingesprongen blok met opdrachten. Wanneer de functie wordt aangeroepen, wordt het lichaam uitgevoerd, waarbij elke parameter de waarde krijgt van het overeenkomstige argument van de functie-oproep. Wanneer we tenslotte de volgende regel in het lichaam van de functie bereiken:</p>
    <p><c>return «waarde»</c></p>
    <p>stopt de functie en geeft <c>«waarde»</c> als resultaatwaarde terug<alert>. </alert>Zoals het <c>abs</c>-voorbeeld laat zien kan het lichaam meerdere <c>return</c>-opdrachten bevatten; alleen de eerste heeft echter effect omdat daarna de functie stopt.</p>
  </section>
  <section xml:id="les-10-s2">
    <title>Voorbeeld: Definitie en aanroep van een functie</title>
    <p>Het volgende programma bevat een definitie van een functie en enkele andere opdrachten die worden uitgevoerd nadat de functie is gedefinieerd.</p>
    <example>
      <p>Kwadraat-functie</p>
      <program interactive="activecode" language="python" label="les-10-vb-2">
        <code>
def kwadraat(x):  # functiedefinitie
    return x*x  # de body bestaat slechts uit een enkele regel

# gebruik de functie, nu deze gedefinieerd is
print(kwadraat(10))        
print(kwadraat(kwadraat(2)))
</code>
      </program>
    </example>
    <p>We zien dat de naam van de functie <c>square</c> is, met slechts één parameter <c>x</c>, en dat het lichaam van de functie bestaat uit één regel: <c>return x*x</c>. Verder, buiten de functie, zijn er twee opdrachten die de functie in totaal drie keer aanroepen.</p>
    <p><ol><li><p>Wanneer de eerste opdracht wordt uitgevoerd, moet Python <c>square(10)</c> evalueren.</p><p><ul><li><p>Python vergelijkt de argumentenlijst <c>(10)</c> met de parameterlijst <c>(x)</c>. Het voert het lichaam van de functie uit, nl.: <c>return x*x</c>, waarbij het zich herinnert dat <c>x</c> gelijk is aan <c>10</c>. Dus geeft <c>x*x</c> het resultaat <c>100</c> en dat wordt afgedrukt.</p></li></ul></p></li><li><p>Wanneer de tweede opdracht wordt uitgevoerd, moet Python <c>square(square(2))</c> evalueren.</p><p><ul><li><p>Het binnenste deel <c>square(2)</c> wordt eerst uitgevoerd. We nemen voor <c>x</c> het getal <c>2</c> en voeren het lichaam van de functie uit. Dat geeft als resultaat <c>4</c>.</p></li><li><p>Dan wordt de buitenste uitdrukking uitgevoerd, en omdat <c>square(2)</c> het getal <c>4</c> heeft teruggegeven, roepen we nu <c>square(4)</c> aan. Dit levert <c>16</c> als resultaatwaarde, en dat wordt afgedrukt.</p></li></ul></p></li></ol></p>
    <p>Zoals gebruikelijk kun je het programma uitproberen in Python3 Visualizer.</p>
  </section>
  <section xml:id="les-10-s3">
    <title>Vier vaak voorkomende fouten</title>
    <p>Een vaak voorkomende fout die je kunt maken bij de definitie van een functie is het vergeten van de <c>return</c>-opdracht.</p>
    <example>
      <p>Fout 1: <c>return</c> vergeten</p>
      <program interactive="activecode" language="python" label="les-10-vb-3">
        <code>
def square(x):
  x*x
print(square(10))
</code>
      </program>
    </example>
    <p>Zoals je kunt zien: wanneer geen <c>return</c>-opdracht in het lichaam voorkomt, geeft de functie <c>None</c> als resultaat. Dus, wanneer je bij een oefening niet het gewenste resultaat krijgt omdat de functie <c>None</c> teruggeeft dan is het probleem vaak dat er bij de uitvoering van de functie geen <c>return</c>-opdracht uitgevoerd is.</p>
    <assemblage>
      <p>Je kunt ook de <c>return</c>-opdracht met reden weglaten. Dit is zinvol als je functie een neveneffect heeft anders dan het geven van een resultaatwaarde:</p>
      <p><alert>Een neveneffect en geen return-opdracht</alert></p>
      <program interactive="activecode" language="python" label="les-10-vb-4">
        <code>
def herhaalPrint(string, herhalingen):   # geen enkele return in deze functie
  for i in range(0, herhalingen):
    print(string)

herhaalPrint('Hello, world!', 3)
</code>
      </program>
    </assemblage>
    <p>Een andere veelvoorkomende fout is vergeten om in te springen, hetgeen resulteert in <c>IndentationError</c>.</p>
    <example>
      <p>Fout 2: vergeten om in te springen</p>
      <program interactive="activecode" language="python" label="les-10-vb-5">
        <code>
def square(x):
return x*x
print(square(10))
</code>
      </program>
    </example>
    <p>Zoals we in les 2 zagen, veroorzaakt het aanroepen van een functie met te weinig of te veel argumenten een fout.</p>
    <example>
      <p>Fout 3: aanroep met te veel argumenten</p>
      <program interactive="activecode" language="python" label="les-10-vb-6">
        <code>
def square(x):
  return x*x
print(square(7, 11))
</code>
      </program>
    </example>
    <p>Tenslotte, wanneer je een tikfout maakt bij de aanroep van een functie, krijg je een foutmelding waarbij gemeld wordt dat de functie niet gedefinieerd is.</p>
    <example>
      <p>Fout 4: verkeerde naam</p>
      <program interactive="activecode" language="python" label="les-10-vb-7">
        <code>
def square(x):
  return x*x
print(skware(10))
</code>
      </program>
    </example>
  </section>
  <section xml:id="les-10-s4">
    <title>Probeer het zelf</title>
    <exercise label="oef-10-cube">
      <title>Tot de derde macht</title>
      <statement>
        <p>Definieer een functie <c>derdeMacht(n)</c>, die één getal <c>n</c> als argument neemt, en <c>n</c> × <c>n</c> × <c>n</c> als resultaat teruggeeft.</p>
      </statement>
      <program interactive="activecode" language="python" label="oef-10-cube-code">
        <code>

</code>
        <tests>
# Verborgen tests voor 10.cube (gegenereerd door tools/harnas.py)
_OPLOSSING = 'def derdeMacht(n):\n    return n * n * n'
_GEVALLEN = [{'aanroep': 'derdeMacht(3)'}, {'aanroep': 'derdeMacht(0)'}, {'aanroep': 'derdeMacht(-2)'}, {'aanroep': 'derdeMacht(10)'}, {'aanroep': 'derdeMacht(1.5)'}]
_VERGELIJK = 'exact'
_TOLERANTIE = 1e-06
_HERGEBRUIK = False
_VERBODEN = []
_VERPLICHT = []
_MAXLENGTE = 0
_MAXWIJZIG = 0
_ORIGINEEL = ''

import re as _re
from unittest.gui import TestCaseGui


class _TeWeinigInvoer(Exception):
    pass


def _voer_uit(_bron, _geval, _student=False):
    _regels = [str(_r) for _r in _geval.get("invoer", [])]
    _uit = []

    def _print(*args, **kw):
        _sep = kw.get("sep", " ")
        _end = kw.get("end", "\n")
        if _sep is None:
            _sep = " "
        if _end is None:
            _end = "\n"
        _uit.append(_sep.join([str(_a) for _a in args]) + _end)

    def _input(prompt=""):
        if len(_regels) == 0:
            raise _TeWeinigInvoer("het programma vraagt meer invoer (input) dan er is")
        return _regels.pop(0)

    _g = {"__name__": "__main__", "print": _print, "input": _input}
    _fout = None
    _waarde = None
    if _student and _HERGEBRUIK:
        # Het programma van de leerling is net al uitgevoerd; gebruik die variabelen.
        try:
            _waarde = eval(_geval["aanroep"], globals())
        except Exception as _e:
            _fout = type(_e).__name__ + ": " + str(_e)
        return "", _waarde, _fout
    try:
        if _geval.get("voorcode"):
            exec(_geval["voorcode"], _g)
        exec(_bron, _g)
        if _geval.get("aanroep"):
            _waarde = eval(_geval["aanroep"], _g)
    except Exception as _e:
        _fout = type(_e).__name__ + ": " + str(_e)
    return "".join(_uit), _waarde, _fout


def _normaal(_s):
    # Spaties aan het eind van een regel tellen niet mee, en het (onzichtbare) DEL-teken ook niet.
    _r = [_x.replace(chr(127), "").rstrip() for _x in _s.split("\n")]
    while len(_r) &gt; 0 and _r[-1] == "":
        _r.pop()
    return "\n".join(_r)


def _getal(_s):
    try:
        return float(_s)
    except Exception:
        return None


def _stukken(_s):
    # Splits in getallen en overige tekens, zodat bv. "-13.333C" -&gt; ["-13.333", "C"].
    return _re.findall(r"[-+]?(?:[0-9]+[.]?[0-9]*|[.][0-9]+)(?:[eE][-+]?[0-9]+)?|[^\s0-9.+-]+|\S", _s)


def _gelijk(_a, _b):
    if _VERGELIJK == "woorden":
        return _a.split() == _b.split()
    if _VERGELIJK == "getallen":
        _wa = _stukken(_a)
        _wb = _stukken(_b)
        if len(_wa) != len(_wb):
            return False
        for _i in range(len(_wa)):
            _x = _getal(_wa[_i])
            _y = _getal(_wb[_i])
            if _x is not None and _y is not None:
                if abs(_x - _y) &gt; _TOLERANTIE * max(1.0, abs(_y)):
                    return False
            elif _wa[_i] != _wb[_i]:
                return False
        return True
    return _normaal(_a) == _normaal(_b)


def _afstand(_a, _b):
    """Bewerkingsafstand (Levenshtein) tussen twee strings."""
    _vorige = list(range(len(_b) + 1))
    for _i in range(len(_a)):
        _huidig = [_i + 1]
        for _j in range(len(_b)):
            _kosten = 0 if _a[_i] == _b[_j] else 1
            _huidig.append(min(_vorige[_j + 1] + 1, _huidig[_j] + 1, _vorige[_j] + _kosten))
        _vorige = _huidig
    return _vorige[-1]


def _zelfde(_a, _b):
    """Vergelijk returnwaarden; bij vergelijk="getallen" mogen kommagetallen iets afwijken."""
    if _VERGELIJK == "getallen" and type(_a) in (int, float) and type(_b) in (int, float) \
            and type(_a) != bool and type(_b) != bool:
        return abs(_a - _b) &lt;= _TOLERANTIE * max(1.0, abs(_b))
    if type(_a) in (list, tuple) and type(_b) == type(_a):
        if len(_a) != len(_b):
            return False
        for _i in range(len(_a)):
            if not _zelfde(_a[_i], _b[_i]):
                return False
        return True
    return type(_a) == type(_b) and _a == _b


def _html(_s):
    return str(_s).replace("&amp;", "&amp;amp;").replace("&lt;", "&amp;lt;").replace("&gt;", "&amp;gt;")


def _blok(_s):
    if _s == "":
        return "&lt;i&gt;(geen uitvoer)&lt;/i&gt;"
    return "&lt;pre style='text-align:left;margin:2px 0'&gt;" + _html(_s) + "&lt;/pre&gt;"


class _Toets(TestCaseGui):
    def test_1_programma(self):
        _bron = self.getEditorText()
        for _tekst, _melding in _VERPLICHT:
            self.assertTrue(_tekst in _bron, _melding)
        for _tekst, _melding in _VERBODEN:
            self.assertFalse(_tekst in _bron, _melding)
        if _MAXLENGTE:
            _lengte = len(_bron.rstrip())
            self.assertTrue(_lengte &lt;= _MAXLENGTE, "Je code is %d tekens lang; het mag hoogstens %d zijn."
                            % (_lengte, _MAXLENGTE))
        if _MAXWIJZIG:
            _d = _afstand(_normaal(_ORIGINEEL), _normaal(_bron))
            self.assertTrue(_d &lt;= _MAXWIJZIG, "Je hebt %d teken(s) veranderd; het mogen er hoogstens %d zijn."
                            % (_d, _MAXWIJZIG))
        _nr = 0
        for _geval in _GEVALLEN:
            _nr += 1
            _verwacht, _vwaarde, _vfout = _voer_uit(_OPLOSSING, _geval)
            _uitvoer, _waarde, _fout = _voer_uit(_bron, _geval, True)
            _info = "&lt;b&gt;Test %d&lt;/b&gt;" % _nr
            if _geval.get("voorcode"):
                _info += "&lt;br&gt;vooraf: &lt;code&gt;" + _html("; ".join(_geval["voorcode"].strip().split("\n"))) + "&lt;/code&gt;"
            if _geval.get("invoer"):
                _info += "&lt;br&gt;invoer:" + _blok("\n".join([str(_r) for _r in _geval["invoer"]]))
            if _geval.get("aanroep"):
                _info += "&lt;br&gt;aanroep: &lt;code&gt;" + _html(_geval["aanroep"]) + "&lt;/code&gt;"
            if _fout is not None and _vfout is None:
                self.appendResult(False, _fout, "", _info + "&lt;br&gt;fout: &lt;code&gt;" + _html(_fout) + "&lt;/code&gt;")
                continue
            _ok = _gelijk(_uitvoer, _verwacht)
            if _geval.get("aanroep"):
                _ok = _ok and _zelfde(_waarde, _vwaarde)
                _jij = _waarde
                _moet = _vwaarde
                if _normaal(_verwacht) != "" or _normaal(_uitvoer) != "":
                    _jij = _normaal(_uitvoer) + " / " + repr(_waarde)
                    _moet = _normaal(_verwacht) + " / " + repr(_vwaarde)
            else:
                _jij = _normaal(_uitvoer)
                _moet = _normaal(_verwacht)
            if not _ok:
                if _geval.get("aanroep"):
                    _info += "&lt;br&gt;verwachte returnwaarde: &lt;code&gt;" + _html(repr(_vwaarde)) + \
                             "&lt;/code&gt;&lt;br&gt;jouw returnwaarde: &lt;code&gt;" + _html(repr(_waarde)) + "&lt;/code&gt;"
                if not _gelijk(_uitvoer, _verwacht) or _normaal(_verwacht) != "":
                    _info += "&lt;br&gt;verwachte uitvoer:" + _blok(_normaal(_verwacht)) + \
                             "jouw uitvoer:" + _blok(_normaal(_uitvoer))
            self.appendResult(_ok, _jij, _moet, _info)


_Toets().main()

</tests>
      </program>
    </exercise>
  </section>
  <section xml:id="les-10-s5">
    <title>Twee of meer argumenten</title>
    <subsection xml:id="les-10-s5-ss0">
      <title>Inleiding</title>
      <image source="images/10-area-triangle.png" width="25%"/>
      <p>De functies hierboven hebben slechts één argument, maar een functie kan een willekeurig aantal argumenten accepteren. Je hebt bijvoorbeeld al <c>input()</c> met 0 argumenten aangeroepen (en je zult een functie definiëren met 0 argumenten in les 15A).</p>
      <p>Hier is een voorbeeld van een functie uit de meetkunde met twee argumenten. Neem aan dat we een functie nodig hebben om de oppervlakte van een driehoek uit te rekenen, waarbij de lengte van de basis en de bijbehorende hoogte zijn gegeven. De formule voor de oppervlakte is</p>
      <p>oppervlakte = basis × hoogte / 2</p>
      <p>Dan kunnen we in Python een functie <c>driehoekOpp</c> als volgt definiëren:</p>
      <example>
        <p>Een functie met twee argumenten</p>
        <program interactive="activecode" language="python" label="les-10-vb-8">
          <code>
def driehoekOpp(basis, hoogte):
  return basis*hoogte/2
print(driehoekOpp(3, 5))
print(driehoekOpp(4, 2))
</code>
        </program>
      </example>
      <p>Zoals je kunt zien, is de enige verandering dat we <c>def «functienaam»(«argument1»)</c> veranderen in <c>def «functienaam»(«argument1», «argument2»)</c>.</p>
      <image source="images/10-peri-rect.png" width="35%"/>
      <p>In de laatste oefening vragen we je om een functie met twee argumenten te schrijven, waarmee je de omtrek van een rechthoek kunt berekenen (de omtrek is de totale lengte van alle zijden). Bij een rechthoek met een gegeven breedte en hoogte, wordt de omtrek gegeven door de volgende formule:</p>
      <p>omtrek = breedte + hoogte + breedte + hoogte</p>
      <exercise label="oef-10-rectangle">
        <title>Rechthoek</title>
        <statement>
          <p>Definieer een functie <c>omtrek(breedte, hoogte)</c>, die de omtrek van een rechthoek berekent.</p>
        </statement>
        <program interactive="activecode" language="python" label="oef-10-rectangle-code">
          <code>

</code>
          <tests>
# Verborgen tests voor 10.rectangle (gegenereerd door tools/harnas.py)
_OPLOSSING = 'def omtrek(breedte, hoogte):\n    return 2 * (breedte + hoogte)'
_GEVALLEN = [{'aanroep': 'omtrek(3, 4)'}, {'aanroep': 'omtrek(10, 10)'}, {'aanroep': 'omtrek(0, 7)'}, {'aanroep': 'omtrek(2.5, 1)'}, {'aanroep': 'omtrek(1, 100)'}]
_VERGELIJK = 'exact'
_TOLERANTIE = 1e-06
_HERGEBRUIK = False
_VERBODEN = []
_VERPLICHT = []
_MAXLENGTE = 0
_MAXWIJZIG = 0
_ORIGINEEL = ''

import re as _re
from unittest.gui import TestCaseGui


class _TeWeinigInvoer(Exception):
    pass


def _voer_uit(_bron, _geval, _student=False):
    _regels = [str(_r) for _r in _geval.get("invoer", [])]
    _uit = []

    def _print(*args, **kw):
        _sep = kw.get("sep", " ")
        _end = kw.get("end", "\n")
        if _sep is None:
            _sep = " "
        if _end is None:
            _end = "\n"
        _uit.append(_sep.join([str(_a) for _a in args]) + _end)

    def _input(prompt=""):
        if len(_regels) == 0:
            raise _TeWeinigInvoer("het programma vraagt meer invoer (input) dan er is")
        return _regels.pop(0)

    _g = {"__name__": "__main__", "print": _print, "input": _input}
    _fout = None
    _waarde = None
    if _student and _HERGEBRUIK:
        # Het programma van de leerling is net al uitgevoerd; gebruik die variabelen.
        try:
            _waarde = eval(_geval["aanroep"], globals())
        except Exception as _e:
            _fout = type(_e).__name__ + ": " + str(_e)
        return "", _waarde, _fout
    try:
        if _geval.get("voorcode"):
            exec(_geval["voorcode"], _g)
        exec(_bron, _g)
        if _geval.get("aanroep"):
            _waarde = eval(_geval["aanroep"], _g)
    except Exception as _e:
        _fout = type(_e).__name__ + ": " + str(_e)
    return "".join(_uit), _waarde, _fout


def _normaal(_s):
    # Spaties aan het eind van een regel tellen niet mee, en het (onzichtbare) DEL-teken ook niet.
    _r = [_x.replace(chr(127), "").rstrip() for _x in _s.split("\n")]
    while len(_r) &gt; 0 and _r[-1] == "":
        _r.pop()
    return "\n".join(_r)


def _getal(_s):
    try:
        return float(_s)
    except Exception:
        return None


def _stukken(_s):
    # Splits in getallen en overige tekens, zodat bv. "-13.333C" -&gt; ["-13.333", "C"].
    return _re.findall(r"[-+]?(?:[0-9]+[.]?[0-9]*|[.][0-9]+)(?:[eE][-+]?[0-9]+)?|[^\s0-9.+-]+|\S", _s)


def _gelijk(_a, _b):
    if _VERGELIJK == "woorden":
        return _a.split() == _b.split()
    if _VERGELIJK == "getallen":
        _wa = _stukken(_a)
        _wb = _stukken(_b)
        if len(_wa) != len(_wb):
            return False
        for _i in range(len(_wa)):
            _x = _getal(_wa[_i])
            _y = _getal(_wb[_i])
            if _x is not None and _y is not None:
                if abs(_x - _y) &gt; _TOLERANTIE * max(1.0, abs(_y)):
                    return False
            elif _wa[_i] != _wb[_i]:
                return False
        return True
    return _normaal(_a) == _normaal(_b)


def _afstand(_a, _b):
    """Bewerkingsafstand (Levenshtein) tussen twee strings."""
    _vorige = list(range(len(_b) + 1))
    for _i in range(len(_a)):
        _huidig = [_i + 1]
        for _j in range(len(_b)):
            _kosten = 0 if _a[_i] == _b[_j] else 1
            _huidig.append(min(_vorige[_j + 1] + 1, _huidig[_j] + 1, _vorige[_j] + _kosten))
        _vorige = _huidig
    return _vorige[-1]


def _zelfde(_a, _b):
    """Vergelijk returnwaarden; bij vergelijk="getallen" mogen kommagetallen iets afwijken."""
    if _VERGELIJK == "getallen" and type(_a) in (int, float) and type(_b) in (int, float) \
            and type(_a) != bool and type(_b) != bool:
        return abs(_a - _b) &lt;= _TOLERANTIE * max(1.0, abs(_b))
    if type(_a) in (list, tuple) and type(_b) == type(_a):
        if len(_a) != len(_b):
            return False
        for _i in range(len(_a)):
            if not _zelfde(_a[_i], _b[_i]):
                return False
        return True
    return type(_a) == type(_b) and _a == _b


def _html(_s):
    return str(_s).replace("&amp;", "&amp;amp;").replace("&lt;", "&amp;lt;").replace("&gt;", "&amp;gt;")


def _blok(_s):
    if _s == "":
        return "&lt;i&gt;(geen uitvoer)&lt;/i&gt;"
    return "&lt;pre style='text-align:left;margin:2px 0'&gt;" + _html(_s) + "&lt;/pre&gt;"


class _Toets(TestCaseGui):
    def test_1_programma(self):
        _bron = self.getEditorText()
        for _tekst, _melding in _VERPLICHT:
            self.assertTrue(_tekst in _bron, _melding)
        for _tekst, _melding in _VERBODEN:
            self.assertFalse(_tekst in _bron, _melding)
        if _MAXLENGTE:
            _lengte = len(_bron.rstrip())
            self.assertTrue(_lengte &lt;= _MAXLENGTE, "Je code is %d tekens lang; het mag hoogstens %d zijn."
                            % (_lengte, _MAXLENGTE))
        if _MAXWIJZIG:
            _d = _afstand(_normaal(_ORIGINEEL), _normaal(_bron))
            self.assertTrue(_d &lt;= _MAXWIJZIG, "Je hebt %d teken(s) veranderd; het mogen er hoogstens %d zijn."
                            % (_d, _MAXWIJZIG))
        _nr = 0
        for _geval in _GEVALLEN:
            _nr += 1
            _verwacht, _vwaarde, _vfout = _voer_uit(_OPLOSSING, _geval)
            _uitvoer, _waarde, _fout = _voer_uit(_bron, _geval, True)
            _info = "&lt;b&gt;Test %d&lt;/b&gt;" % _nr
            if _geval.get("voorcode"):
                _info += "&lt;br&gt;vooraf: &lt;code&gt;" + _html("; ".join(_geval["voorcode"].strip().split("\n"))) + "&lt;/code&gt;"
            if _geval.get("invoer"):
                _info += "&lt;br&gt;invoer:" + _blok("\n".join([str(_r) for _r in _geval["invoer"]]))
            if _geval.get("aanroep"):
                _info += "&lt;br&gt;aanroep: &lt;code&gt;" + _html(_geval["aanroep"]) + "&lt;/code&gt;"
            if _fout is not None and _vfout is None:
                self.appendResult(False, _fout, "", _info + "&lt;br&gt;fout: &lt;code&gt;" + _html(_fout) + "&lt;/code&gt;")
                continue
            _ok = _gelijk(_uitvoer, _verwacht)
            if _geval.get("aanroep"):
                _ok = _ok and _zelfde(_waarde, _vwaarde)
                _jij = _waarde
                _moet = _vwaarde
                if _normaal(_verwacht) != "" or _normaal(_uitvoer) != "":
                    _jij = _normaal(_uitvoer) + " / " + repr(_waarde)
                    _moet = _normaal(_verwacht) + " / " + repr(_vwaarde)
            else:
                _jij = _normaal(_uitvoer)
                _moet = _normaal(_verwacht)
            if not _ok:
                if _geval.get("aanroep"):
                    _info += "&lt;br&gt;verwachte returnwaarde: &lt;code&gt;" + _html(repr(_vwaarde)) + \
                             "&lt;/code&gt;&lt;br&gt;jouw returnwaarde: &lt;code&gt;" + _html(repr(_waarde)) + "&lt;/code&gt;"
                if not _gelijk(_uitvoer, _verwacht) or _normaal(_verwacht) != "":
                    _info += "&lt;br&gt;verwachte uitvoer:" + _blok(_normaal(_verwacht)) + \
                             "jouw uitvoer:" + _blok(_normaal(_uitvoer))
            self.appendResult(_ok, _jij, _moet, _info)


_Toets().main()

</tests>
        </program>
      </exercise>
    </subsection>
    <subsection xml:id="les-10-s5-ss1">
      <title>Functies die functies aanroepen</title>
      <p>Functies zijn de bouwstenen van goed geschreven grote programma's: je kunt dezelfde taak twee keer uitvoeren zonder dezelfde code twee keer te schrijven, en je kunt oplossingen van veelvoorkomende taken opnieuw gebruiken. Hier is een voorbeeld van het schrijven van een functie die gebruikt wordt in een andere functie.</p>
      <example>
        <p><em>Kun je voorspellen wat deze functie doet voordat je het programma uitvoert?</em></p>
        <p>Herinner je uit les 4 dat de vermenigvuldiging van strings met een geheel getal inhoudt dat de string dit aantal keren herhaald wordt. Zo is bijvoorbeeld <c>"tar"*2 gelijk aan "tartar"</c>.</p>
        <program interactive="activecode" language="python" label="les-10-vb-9">
          <code>
def inner(string, width):
  return string + " " * (width - len(string))

def outer(string1, string2):
  w = max(len(string1), len(string2))
  print("*" * (w + 4))
  print("* " + inner(string1, w) + " *")
  print("* " + inner(string2, w) + " *")
  print("*" * (w + 4))

outer("Box number one", "is very fun")
outer("Box #2", "(not False) is " + str(not False))
</code>
        </program>
      </example>
      <p>Je bent nu klaar voor de volgende les!</p>
    </subsection>
  </section>
</chapter>
